Ultrastructural sublaminar-specific diversity of excitatory synaptic boutons in layer 1 of the adult human temporal lobe neocortex

Author response:

The following is the authors’ response to the previous reviews.

Public Reviews:

Reviewer #1 (Public review):

Summary:

The authors investigated the anatomical features of the synaptic boutons in layer 1 of the human temporal neocortex. They examined the size of each synapse, the macular or perforated appearance, the size of the synaptic active zone, the number and volume of the mitochondria, and the number of synaptic and dense core vesicles, also differentiating between the readily releasable, the recycling, and the resting pool of synaptic vesicles. The coverage of the synapse by astrocytic processes was also assessed, and all the above parameters were compared to other layers of the human temporal neocortex. The authors conclude that the subcellular morphology of the layer 1 synapses are suitable for the functions of the neocortical layer, i.e. the synaptic integration within the cortical column. The low glial coverage of the synapses might allow increased glutamate spillover from the synapses, enhancing synaptic crosstalk within this cortical layer.

Strengths:

The strengths of this paper are the abundant and very precious data about the fine structure of the human neocortical layer 1. Quantitative electron microscopy data (especially that derived from the human brain) are very valuable since this is a highly time- and energy-consuming work. The techniques used to obtain the data, as well as the analyses and the statistics performed by the authors are all solid, strengthen this manuscript, and mainly support the conclusions drawn in the discussion.

We would like to thank reviewer#1 for his very positive comments on our manuscript stating that such data about the fine structure of the human neocortex are are highly relevant.

Weaknesses:

There are several weaknesses in this work. First, the authors should check and review extensively for improvements to the use of English. Second, several additional analyses performed on the existing data could substantially elevate the value of the data presented. Much more information could be gained from the existing data about the functions of the investigated layer, of the cortical column, and about the information processing of the human neocortex. Third, several methodological concerns weaken the conclusions drawn from the results.

We would like to thank the reviewer for his critical and thus helpful comments on our manuscript. We took the first comment of the reviewer concerning the English and have thus improved our manuscript by rephrasing and shortening sentences. Secondly, according to the reviewer several additional analyses should be performed on the existing data, which could substantially elevate the value of the data presented. We will implement some of the suggestions in the improved version of the manuscript where appropriate. We will address a more detailed answer to the reviewer’s queries in her/his suggestions to the authors (see below). However, the reviewer states himself: “The techniques used to obtain the data, as well as the analyses and the statistics performed by the authors are all solid, strengthen this manuscript, and mainly support the conclusions drawn in the discussion”.

Reviewer #2 (Public review):

Summary:

The study of Rollenhagen et al. examines the ultrastructural features of Layer 1 of the human temporal cortex. The tissue was derived from drug-resistant epileptic patients undergoing surgery, and was selected as far as possible from the epilepsy focus, and as such considered to be non-epileptic. The analyses included 4 patients with different ages, sex, medication, and onset of epilepsy. The manuscript is a follow-on study with 3 previous publications from the same authors on different layers of the temporal cortex:

Layer 4 - Yakoubi et al 2019 eLife

Layer 5 - Yakoubi et al 2019 Cerebral Cortex

Layer 6 - Schmuhl-Giesen et al 2022 Cerebral Cortex.

They find, that the L1 synaptic boutons mainly have a single active zone, a very large pool of synaptic vesicles, and are mostly devoid of astrocytic coverage.

Strengths:

The manuscript is well-written and easy to read. The Results section gives a detailed set of figures showing many morphological parameters of synaptic boutons and glial elements. The authors provide comparative data of all the layers examined by them so far in the Discussion. Given that anatomical data in the human brain are still very limited, the current manuscript has substantial relevance. The work appears to be generally well done, the EM and EM tomography images are of very good quality. The analysis is clear and precise.

We would like to thank the reviewer for his very positive evaluation of our paper and the comments that such data have a substantial relevance, in particular in the human neocortex. In contrast to reviewer#1, this reviewer’s opinion is that the manuscript is well written and easy to read.

Weaknesses:

One of the main findings of this paper is that "low degree of astrocytic coverage of L1 SBs suggests that glutamate spillover and as a consequence synaptic cross-talk may occur at the majority of synaptic complexes in L1". However, the authors only quantified the volume ratio of astrocytes in all 6 layers, which is not necessarily the same as the glial coverage of synapses. In order to strengthen this statement, the authors could provide 3D data (that they have from the aligned serial sections) detailing the percentage of synapses that have glial processes in close proximity to the synaptic cleft, that would prevent spillover.

We agree with the reviewer that we only quantified the volume ratio of the astrocytic coverage but not necessarily the percentage of synapses that may or not contribute to the formation of the ‘tripartite’ synapse. As suggested, we will re-analyze our material with respect to the percentage of coverage for individual synaptic boutons in each layer and will implement the results in the improved version of the manuscript. However, since this is a completely new analysis that is time-consuming we would like to ask the reviewer for additional time to perform this task.

A specific statement is missing on whether only glutamatergic boutons were analyzed in this MS, or GABAergic boutons were also included. There is a statement, that they can be distinguished from glutamatergic ones, but it would be useful to state it clearly in the Abstract, Results, and Methods section what sort of boutons were analyzed. Also, what is the percentage of those boutons from the total bouton population in L1?

We would like to thank the reviewer for this comment. Although our title clearly states, we focused on quantitative 3D-models of excitatory synaptic boutons, we will point out that more clearly in the Methods and Result chapters. Our data support recent findings by others (see for example Cano-Astorga et al. 2023, 2024; Shapson-Coe et al. 2024) that have evaluated the ratio between excitatory vs. inhibitory synaptic boutons in the temporal lobe neocortex, the same area as in our study, which was between 10-15% inhibitory terminals but with a significant layer and region specific difference. We will include the excitatory vs. inhibitory ratio and the corresponding citations in the Results section.

Synaptic vesicle diameter (that has been established to be ~40nm independent of species) can properly be measured with EM tomography only, as it provides the possibility to find the largest diameter of every given vesicle. Measuring it in 50 nm thick sections results in underestimation (just like here the values are ~25 nm) as the measured diameter will be smaller than the true diameter if the vesicle is not cut in the middle, (which is the least probable scenario). The authors have the EM tomography data set for measuring the vesicle diameter properly.

We partially disagree with the reviewer on this point. Using high-resolution transmission electron microscopy, we measured the distance from the outer-to-outer membrane only on those synaptic vesicles that were round in shape with a clear ring-like structure to avoid double counts and discarded all those that were only partially cut according to criteria developed by Abercrombie (1946) and Boissonnat (1988). We assumed that within a 55±5 nm thick ultrathin section (silver to gray interference contrast) all clear-ring-like vesicles were distributed in this section assuming a vesicle diameter between 25 to 40nm. For large DCVs, double-counts were excluded by careful examination of adjacent images and were only counted in the image where they appeared largest.

In addition, we have measured synaptic vesicles using TEM tomography and came to similar results. We will address this in Material and Methods that both methods were used.

It is a bit misleading to call vesicle populations at certain arbitrary distances from the presynaptic active zone as readily releasable pool, recycling pool, and resting pool, as these are functional categories, and cannot directly be translated to vesicles at certain distances. Indeed, it is debated whether the morphologically docked vesicles are the ones, that are readily releasable, as further molecular steps, such as proper priming are also a prerequisite for release.

We thank the reviewer for this comment. However, nobody before us tried to define a morphological correlate for the three functionally defined pools of synaptic vesicles since synaptic vesicles normally are distributed over the entire nerve terminal. As already mentioned above, after long and thorough discussions with Profs. Bill Betz, Chuck Stevens, Thomas Schikorski and other experts in this field we tried to define the readily releasable (RRP), recycling (RP) and resting pools by measuring the distance of each synaptic vesicle to the presynaptic density (PreAZ). Using distance as a criterion, we defined the RRP including all vesicles that were located within a distance (perimeter) of 10 to 20 nm from the PreAZ that is less than an average vesicle diameter (between 25 to 40 nm). The RP was defined as vesicles within a distance of 60-200 nm away, still quite close but also rapidly available on demand and the remaining ones beyond 200 nm were suggested to belong to the resting pool. This concept was developed for our first publication (Sätzler et al. 2002) and this approximation since then is very much acknowledged by scientist working in the field of synaptic neuroscience and computational neuroscientist. We were asked by several labs worldwide whether they can use our data of the perimeter analysis for modeling. We agree that our definition of the three pools can be seen as arbitrary but we never claimed that our approach is the truth but nothing as the truth. Concerning the debate whether only docked vesicles or also those very close the PreAZ should constitute the RRP we have a paper in preparation using our perimeter analysis, EM tomography and simulations trying to clarify this debate. Our preliminary results suggest that the size of the RRP should be reconsidered.

Tissue shrinkage due to aldehyde fixation is a well-documented phenomenon that needs compensation when dealing with density values. The authors cite Korogod et al 2015 - which actually draws attention to the problem comparing aldehyde fixed and non-fixed tissue, still the data is non-compensated in the manuscript. Since all the previous publications from this lab are based on aldehyde fixed non-compensated data, and for this sake, this dataset should be kept as it is for comparative purposes, it would be important to provide a scaling factor applicable to be able to compare these data to other publications.

We thank the reviewer for his suggestion. However, for several reasons we did not correct for shrinkage caused by aldehyde fixation. There are papers by Eyre et al. (2007) and the mentioned paper by Korogod et al. 2015 that have demonstrated that cryo-fixation reveals larger numbers of docked synaptic vesicles, a smaller glial volume, and a less intimate glial coverage of synapses and blood vessels compared to chemical fixation. Other structural subelements such as active zone size and shape and the total number of synaptic vesicles remained unaffected. In two further publications Zhao et al. (2012a, b) investigating hippocampal mossy fiber boutons using cryo-fixation and substitutions came to similar results with respect to bouton and active zone size and number and diameter of synaptic vesicles compared to aldehyde-fixation as described by Rollenhagen et al. 2007 for the same nerve terminal. This was one of the reasons not correcting for shrinkage. In addition, all cited papers state that chemical fixation in general provides a much better ultrastructural preservation of tissue samples when compared with cryo-fixation and substitution where optimal preservation is only regional within a block of tissue and therefore less suitable for large-scale ultrastructural analyses as we performed.

Reviewer #3 (Public review):

Summary:

Rollenhagen et al. offer a detailed description of layer 1 of the human neocortex. They use electron microscopy to assess the morphological parameters of presynaptic terminals, active zones, vesicle density/distribution, mitochondrial morphology, and astrocytic coverage. The data is collected from tissue from four patients undergoing epilepsy surgery. As the epileptic focus was localized in all patients to the hippocampus, the tissue examined in this manuscript is considered non-epileptic (access) tissue.

Strengths:

The quality of the electron microscopic images is very high, and the data is analyzed carefully. Data from human tissue is always precious and the authors here provide a detailed analysis using adequate approaches, and the data is clearly presented.

We are very thankful to the reviewer upon his very positive comments about our data analysis and presentation.

Weaknesses:

The study provides only morphological details, these can be useful in the future when combined with functional assessments or computational approaches. The authors emphasize the importance of their findings on astrocytic coverage and suggest important implications for glutamate spillover. However, the percentage of synapses that form tripartite synapses has not been quantified, the authors' functional claims are based solely on volumetric fraction measurements.

We thank the reviewer for his critical comments on our findings concerning the layer-specific astrocytic coverage as also suggested by reviewer#2. As already stated above we will analyze the astrocytic coverage and the layer-specific percentage of astrocytic contribution to the ‘tripartite’ synapse in more detail. We are, however, a bit puzzled about the comment that structural anatomists usually receive that our study only provides morphological details. Our thorough analysis of structural and synaptic parameters of synaptic boutons underlie and might even predict the function of synaptic boutons in a given microcircuit or network and will thus very much improve our understanding and knowledge about the functional properties of these structures, in particular in the human brain where such studies are still quite rare. The main goal of our studies in the human neocortex was the quantitative morphology of synaptic boutons and thus the synaptic organization of the cortical column, layer by layer which to our knowledge is the first such detailed study undertaken in the human brain. Our efforts have set a golden standard in the analysis of synaptic boutons embedded in different microcircuits und is meanwhile internationally very well accepted.

The distinction between excitatory and inhibitory synapses is not clear, they should be analyzed separately.

As already stated above in response to reviewer#1 our study focused on excitatory synaptic boutons since they represent the majority of synapses. However, in the improved version of our manuscript in the Material and Method section we included a paragraph with structural criteria to distinguish excitatory from inhibitory terminals (see also our comment to reviewer#1 concerning this point) including appropriate citations.

The text connects functional and morphological characteristics in a very direct way. For example, connecting plasticity to any measurement the authors present would be rather difficult without any additional functional experiments. References to various vesicle pools based on the location of the vesicles are also more complex than suggested in the manuscript. The text should better reflect the limitations of the conclusions that can be drawn from the authors' data.

We thank the reviewer for this comment. However, it has been shown by meanwhile numerous publications that the shape and size of the active zone together with the pool of synaptic vesicles and the astrocytic coverage critically determines synaptic transmission and synaptic strength, but can also contribute to the modulation of synaptic plasticity (see also citations within the text). It has been shown that synaptic boutons can switch upon certain stimulation conditions to different modes of release (uni- vs. multiquantal, uni- vs multivesicular release) and from asynchronous to synchronous release leading also to the modulation of synaptic short- and long-term plasticity. To the second comment: When we started with our first paper about the Calyx of Held – principal neuron synapse in the MNTB (Sätzler et al. 2002) we tried to define a morphological correlate for the three functionally defined pools. As already mentioned above in our reply to the other two reviewers, this is rather difficult since synaptic vesicles are normally distributed over the entire nerve terminal. After long and thorough discussions with Bill Betz, Chuck Stevens and other leading scientist in the field of synaptic neuroscience, we together with Bert Sakmann tried to define a morphological correlate for the functionally defined pools using a perimeter analysis. We defined the readily releasable pool as vesicles 10 to 20 nm away from the presynaptic active zone, the recycling pool as those in 60-200 nm distance and the remaining as those belonging to the resting pool. However, it has been shown by capacitance measurements (see for example Hallermann et al 2003), FM1-43 investigations (see for example Henkel et al. 1996) and high-resolution electron microscopy (see for example Schikorski and Stevens 2001; Schikorski 2014) that our estimate of the RRP nearly perfectly matches with the functionally defined pools at hippocampal and cortical synapses (Silver et al. 2003). In addition, in one of our own papers (Rollenhagen et al. 2018) we also estimated the RP functionally from trains of EPSPs using an exponential fit analysis and came to similar results upon its size using the perimeter analysis.

Of course, as stated by the reviewer the scenario could be more complex, using other criteria but we never claimed that our morphologically defined pools are the truth but nothing as the truth but we believe it offers a quite good approximation.

Recommendations for the authors:

Reviewer #1 (Recommendations for the authors):

Abstract:

Avoid the numerous abbreviations in the abstract. The paragraph describing the results obtained in this study is too short. Include more results, such as the size of the active zone, the proportion of perforated synapses, the ratio of synapses terminating on dendrites/spines, the percentage of volume occupied by mitochondria, etc. In the last paragraph, compare the layer-specific data to other layers of the neocortex before writing the concluding sentence.

To meet the word limits of the abstract (150 words) defined by eLife we had to use abbreviations. We followed the suggestions by the reviewer and expanded our abstract by adding the proportion of macular vs. perforated active zone and the percentage of mitochondria within an SB. However, we did not include the comparison of structural parameters in the Abstract since this is discussed thoroughly in the MS at other places (see Results and Discussion).

Results:

First of all, wonderful data! Lots of work, very valuable quantitative electron microscopy results.

Main concerns:

Adding several analyses would give much more information about the cortical synaptic organization. It would be very useful to differentiate between excitatory and inhibitory terminals (and give their ratio) and include this information in all different analyses, such as in the SV number, SV pool analysis, mitochondrion analysis, etc., that would give functional information as well. You have all the data for this, and you know how to differentiate between inhibitory and excitatory synapses, it can be done. We could see the possible morphological differences between excitatory and inhibitory synapses (maybe one is larger/has more SVs, etc. than the other). Based on these possible differences conclusions could be drawn about functional hypotheses, such as one or the other is more efficient in inducing postsynaptic potentials, excitation or inhibition is more pronounced in layer 1, etc. Furthermore, looking at the ratio of perforated synapses, we could gain information about the formation of new synapses. Maybe there is a difference between excitatory and inhibitory circuits in this point of view.

To the first point: Since our focus was on excitatory synaptic boutons as already stated in the title we have not analyzed inhibitory SBs. To do so, we have to re-analyze our complete data which is time-consuming and an additional workload. However, we can give a ratio excitatory vs. inhibitory synaptic boutons which was between 10-15% but with layer-specific differences. Our finding are in good agreement with a recent publication in Science by the Lichtman group (Shapson-Coe et al. 2024) and work by the DeFelipe group (Cano-Astorga et al. 2023, 2024) estimating the number of inhibitory boutons in different layers of the temporal lobe neocortex as we did by 10-15%. We included a small paragraph about inhibitory synapses, their percentage and included the citations in our Results section. Concerning the ratio between macular, non-perforated vs. perforated active zones we stated the majority of synaptic boutons were of the macular, non-perforated type (~75%; see improved version of the MS). If perforated, this was found predominantly on the postsynaptic site, but quite rare in L1 SBs. Since GABAergic terminals had only a small or no clearly visible PSD this would be hard to look at.

To the last point, it has been demonstrated that the number of dense core vesicles and their fusion with the presynaptic density could be a critical factor in the build-up of the active zone. In addition, the findings of the Geinismann group suggesting that perforated synapses are more efficient than non-perforated ones is nowadays very controversially discussed since other factors such as size of the active zone (see for example Matz et al. 2010; Holderith et al. 2012) and the astrocytic coverage contribute to synaptic efficacy and strength.

Related to this topic: although in the case of rat CA1 pyramidal cells all inhibitory synapses terminated on dendritic shafts (Megias et al., Neuroscience 2001), please be aware that both excitatory and inhibitory synapses can terminate on both dendritic shafts and spines in humans (inhibitory synapses are though rare on spines, usually less than 10%, but they do exist, see for example Wittner et al, Neuroscience, 2001). Please, define the excitatory/inhibitory nature of the synapses based on morphological features (not on their postsynaptic target), i.e., flattened vesicles and thin postsynaptic density for GABAergic synapses, whereas larger, round vesicles and thick postsynaptic density for glutamatergic synapses. Anyway, the ratio of excitatory and inhibitory synapses on dendrites and spines in the two sublamina would also give useful information about the synaptic organization of the human neocortical layer 1.

We are aware that not all terminals targeting on spines are excitatory, in turn it has been shown that not all terminals on shafts were inhibitory as long thought (Silver et al. 2003). However, as stated by the reviewer their abundancy on spines is rather low. At the moment it is rather unclear which functional impact inhibitory terminals on spines have, despite a local inhibition (see for example Kubota et al. eLife 2015), and thus their role is rather speculative since excitatory synapses are the predominant class on dendritic spines. As already stated above the ratio of excitatory vs. inhibitory terminals is between 10-15% and not significantly different between the two sublaminae. We are willing to add this in the results section (see in the improved version of the manuscript).

(2) About the glial coverage: Please, specify how glial elements were determined. What were the morphological features specific to astroglial processes? In Figure 5, how could we know whether the glial element marked by green is not a spine neck? The lack of morphological features specific to glial processes makes this analysis weak. The most accurate would be to make it with the aid of GFAP staining. I know this is not possible with your existing data, but at least, provide information on how glial processes were identified.

We used the criteria first described by Peters et al. (1991) and Ventura and Harris (1999) identifying astrocytic profiles by their irregular stellate shape, relatively clear cytoplasm, numerous glycogen granules and bundles of intermediate filaments. After more than 20 years of structural investigations, we hope that the reviewers will believe us that we can identify astrocytic processes at the high-resolution TEM level. In some of our publications (Rollenhagen et al. 2007; 2015; 2018; Yakoubi et al. 2019a) we have used glutamine synthetase pre-embedding immunhistochemistry to identify astrocytic processes, but a disadvantage of this method is the reduction of the ultrastructural preservation of the tissue. We have included the criteria to identify astrocytic processes of glial coverage in our manuscript together with the two citations (see improved version of the manuscript).

(3) The authors state that the total number of SVs was very variable. How was the distribution of the number of SVs? Homogenous distribution suggests that different types of synapses cannot be distinguished based on their morphological features, whereas distribution with more than one peak would suggest that different types of synapses are present in L1, and that they can be differentiated by their appearance (number of SVs, for example). This might be also related to the type of synapse (i.e., excitatory or inhibitory). The same applies to the number of RP and resting pool SVs.

To look for differences in structural and synaptic parameters that can further classify synaptic boutons we have performed a hierarchical cluster and multivariance analysis. However, it turned out that according to structural and functional parameters no further classification into subtypes could be done.

(4) The authors should check and review extensively for improvements to the use of English. The Results and Discussion sections contain many sentences which are not easy to understand. They have either a too complicated structure, or they are incomplete and hard to follow. Few examples: "The RRP/PreAZ at p20 nm criterium was on average 19.05 {plus minus} 17.23 SVs (L1a: 25.04 {plus minus} 21.09 SVs and L1b: 13.07 {plus minus} 13.87SVs) and thus nearly 2-fold larger for L1a." If you take out the parenthesis, the sentence has no meaning. "The majority of SBs in L1 of the human TLN had a single at most three AZs that could be of the non-perforated macular or perforated type comparable with results for other layers in the human TLN but by ~1.5-fold larger than in rodent and non-human primates." Rephrase these types of sentences, please.

We partially agree with the reviewer. We have improved our manuscript by rephrasing and shortening sentences.

Other suggestions:

(1) Put the synaptic density part after the description of the neuronal and synaptic composition part, it is more logical this way (i.e., first qualitative description, the distinction between sublayers, then quantitative data). Please write down in the description of the neuronal and synaptic composition part how L1a and L1b were differentiated (see also my comment on Figure 1).

We agree with the reviewer and did the change according to the suggestion. For a better understanding, we have also expanded the neuronal and synaptic description of the two sublaminae in L1.

(2) Introduce a list of abbreviations at the beginning, that would help.

It is quite unusual to provide a list of abbreviations in eLife. However, when used first the full meaning of the abbreviations is now given.

(3) What is cleft width? Usually, it refers to the distance between the pre- and the postsynaptic membrane, but here, I think it refers to the size (diameter) of the active zone. Please, clarify in the Result section (as it appears earlier than the Methods section, where it is explained). I would probably use the expression "synaptic cleft size" instead of "synaptic cleft width" to avoid misunderstanding.

We thank the reviewer for the suggestion and used synaptic cleft size for better clarity and have transferred the sentence from the Material and Methods to the Results section.

(4) The description of the different SVs (RRP, RP, etc.) is not clear in lines 236-242. What does it mean, that RRP vesicles are located {less than or equal to}10 nm and {less than or equal to}20 nm from the active zone? Explain, why the two different distance criteria were used. Furthermore, how were the vesicles located at p20-p60 defined? Why were these vesicles not considered in the determination of the different pools?

As stated in the public review to the reviewers concern we have tried to define a morphological correlate to the three functionally defined pools. After thorough discussions, with leading scientists in the field of synaptic neuroscience we have decided to use the distance of individual vesicles from the PreAZ and sort vesicles upon these criteria. One can argue that this approach is random, however, these distance criteria were described by Rizzoli and Betz (2004, 2005) and Denker and Rizzoli (2010). As also stated in the public review there is still a controversial discussion whether only docked or omega-shaped SVs constitute the RRP. We decided that also those very close within 10 and 20 nm away from the PreAZ, which is less than a SV diameter may also contribute to the RRP since it was shown that SVs are quite mobile.

(5) Please, explain how the number of docked vesicles can be 3x larger in L1b, than the number of vesicles located at p10? Docked vesicles are the closest (with the membrane touching the PreAZ)... if this comes from the fact that another pool of boutons was used for the EM tomography analysis, then the entire pool of boutons analyzed, then it means that the selection of boutons for the EM tomography is highly biased. This also implies that EM tomography data are most probably not valid for the entire L1b. The difference might also come from the different ratios of dendrite/spine synapses included in the two different analyses. In this case, it would be helpful to distinguish between synapses terminating on dendrites/spines and analyse them separately (same as for inhibitory/excitatory, which is not exactly the same as dendrite/spine!). Different n numbers of synapses are given in the text (n=25, 25, 25 25) and in Table 2 (n=91, 98, 87, and 84) for the analysis of the docked vesicles, please, correct this.

This is a correct value and thus there is a nearly 3-fold difference. The TEM tomography was carried out on the same blocks that have been used for our 3D-volume reconstructions. To carry out TEM tomography we had to use thicker sections (250 nm) to look for complete SBs as we also did in our serial sections, but of course, we could not quantify the same SBs. The completeness of SBs was one of our main criteria to reconstruct structural and synaptic parameters. The second was that the synaptic cleft was cut perpendicular. Only SBs that met these criteria were chosen for further quantitative analysis. In this respect we are of course biased in both methods.

Secondly, as already stated we did not quantify inhibitory terminals in serial sections. However, we did not find significant differences between shaft vs. spine synapses.

Finally, in Table 2 the total number of ‘docked’ SVs is given analyzed from the total number of SBs analyzed.

Discussion:

Please include the recent findings of human L1 neurons, including the "rosehip" cells in the L1 neuronal network, see Boldog et al., Nat Neurosci 2018. It would be also useful to consider in the discussion the human-specific cortical synchrony and integration phenomena derived from in vitro data (Mansvelder, Lein, Tamas, Wittner, Larkum, Huberfeld labs, etc.), and how the synaptic morphology can be related to these.

We thank the reviewer and include the reference in our chapter functional significance.

Figures and Tables:

Figure 1: In the legend, it is written that CR cells are marked by an asterisk, but on the figure it is marked by arrowheads. H: I would put the dashed line slightly lower, just above the two neuronal cell bodies. Now it looks like in the middle of the astrocytic layer. One of the asterisks marking the CR cell is not above the nucleus of that cell. I: the gabaergic neuron is outside of the framed area. I would delete the frame, anyway, the arrowheads and the asterisk are enough to show what the authors want to show.

We have changed the Figure according to the suggestions raised by the reviewer.

Figure 3: The transparent yellow is not visible. It is a bit disturbing that the contours of the boutons are not visible, I would make the transparent yellow stronger (less transparent). The SVs in green/magenta will be still visible.

We wanted to highlight the internal subelements of SBs and thus made the covering transparent but we think it is still visible.

Figure 6C: The data concerning other layers than L1 are most probably taken from other publications of the research group. One is cited (for L6), but not the others. Please correct this, or if not, then write this in the Results and Methods.

We changed the citation in the improved version of the manuscript. We overlooked that the values for L4 and L5 were already published in Schmuhl-Giesen et al. 2022.

Table 1: What does central and lateral cleft width mean in Table 1? Furthermore, please, give the name for abbreviations CV and IQR in Tables 1 and 2.

The measurements of the synaptic cleft are now described in detail in the Results section. We now have given the full names for CV and IQR in the legends of tables 1 and 2.

Supplemental Figures 1 and 2: Why Hu01 and Hu02 are twice? What is the difference? Based on the figure legend, it is L1a and L1b? If yes, please, indicate on the figure or in the legend.

Supplemental Table 1: What is TLE in the case of Hu_04? If it is temporal lobe epilepsy, then why age at epilepsy onset is missing?

Yes, Hu01 and Hu02 were selected for both L1a and L1b in separate serial sections preparations each. We indicated this now in the figure legend. Concerning Hu_04, unfortunately we do not have any further information about the medical background of the patient.

Supplemental Table 1 (Patient table), that there are many abbreviations explained which do not appear in the table (lBAZ: Brivaracetam CBZ: Carbamazepine; CLB: Clobazam; ESL: Eslicarbazepin; GGL: Ganglioglioma, etc.), please check and correct.

We have removed the unnecessary abbreviations.

Other minor suggestions:

What is Pr? Please, give the name a first appearance (line 368).

We explained Pr (release probability) when used for the first time.

Give the name for t-LDT, please (lines 442-443).

We explained t-LTD (timing-dependent long-term depression) when used for the first time.

Typo in line 169: DCW instead of DCV (dense core vesicle), DCV is used in the figure legends.

We changed DCW to DCV.

Typo in line 190: Yokoubi instead of Yakoubi (reference).

We changed Yokoubi to Yakoubi.

Typo in line 237: Rizzoloi instead of Rizzoli (reference).

We changed Rizzoloi to Rizzoli.

Line 229-230: One reference is not inserted properly - Piccolo and Bassoon.

The reference of Schoch and Gundelfinger and Murkherjee to the build-up of the active zone and the role of DCV containing Piccolo and Bassoon are properly cited in the text.

Typo in line 398: exit instead of exist.

Corrected

Typo in line 700: Reynolds (1063) instead of 1963.

Corrected

Reviewer #2 (Recommendations for the authors):

Abstract:

The last sentence seems far-fetched, and unrelated to the manuscript. How mostly single active zone boutons can "mediate, integrate and synchronize contextual and cross-modal information, enabling flexible and state-dependent processing of feedforward sensory inputs from other layers of the cortical column"? Which of the anatomical findings of the manuscript led to these conclusions?

According to the review by Schuman et al. (2021) layer 1 is regarded as a layer that mediate, integrate and synchronize contextual and cross-modal information, enabling flexible and state-dependent processing of feedforward sensory inputs from other layers of the cortical column to which the structural quantitative 3D- models of SBs contribute since they are an integral element connecting neurons and building networks.

I am also puzzled by the authors' statement in more than one place of the manuscript that "L1a can be characterized as a predominantly astrocytic sublamina". If the L1 contains the lowest measured volume ratio of glial processes (Figure 6), then this description does not seem to hold. Please rephrase.

The reviewer is right and we rephrased the sentences for more clarity in the improved version of our manuscript.

Results:

The authors find large inter-patient variability in the synapse density at L1, which raises the issue of what were the criteria to include certain patients in the analyses. Apparently, these are different from the ones analysed in their previous papers, and all the provided parameters were different (sex, age, medication, onset of epilepsy), and any of them can result in altered synapse density.

First, we have not used all patients for this study. Secondly, it was not possible to use all patients for all six layers.

It would be useful to add a panel for Figure 1 with synapse density across the different layers, as they provide this data in the Discussion.

We implemented a Supplementary Table 1 with the synaptic density values over all layers compared in the Discussion.

I cannot find Source Data 1 in the manuscript although it is referred to in more than 1 place (e.g. page 5 line 100).

Source data were uploaded when our manuscript was submitted directly to eLife as Supplemental Material. However, as stated by bioRxiv ‘any Supplemental Materials associated with this manuscript have not been transferred to bioRxiv to avoid the posting of potentially sensitive information’ all source data have not been uploaded to the preprint server.

Page 5 line 100 the correct value is 7.3*107 or rather 108?

We corrected the value in the improved version of the MS.

It would be nice to put the synapse density values into context by comparing them to e.g. mouse, rat, or monkey data.

Since we are working on the human temporal lobe neocortex we avoided to compare those data with those estimated in experimental animals. In addition as discussed by DeFelipe et al. (1999) different methods were used to quantify synaptic density in experimental animals so these results are difficult to compare.

Page 5 Line 117 CR-cells stands for Cayal-Retzius cells?

CR-cells is the abbreviation for Cajal-Retzius cells.

Page 6 Line 146 repeated sentence.

We deleted the repeated sentence.

Page 7 Line 154 "file-scale TEM" ??

We replaced file-scale by fine-scale.

Page 7 Line 164 "GABAergic synapses identified by the smaller more spherical SVs". With this fixation condition, GABAergic vesicles are more ovoid than glutamatergic ones. What were the criteria to distinguish them?

To our knowledge in meanwhile numerous publications using the same fixation inhibitory terminals contain more spherical and smaller and not roundish synaptic vesicles and showed no clear prominent PSDs as described in our paper. We have addressed that more clearly in the results section of the improved version of the MS.

Page 8 line 197 "The majority (~98%) of SBs in L1a and L1b had only a single (Figures 2C-E, 3A-C, E) at most two or three AZs" is in striking contrast with the other statement from page 7 Line 163 "Numerous SBs in both sublaminae were seen to establish either two or three synaptic contacts on the same spine or dendrite". Which of these statements is valid? Please provide exact quantification for this statement and decide which one is true.

It is true that the majority of synaptic boutons had a single active zone. However, for example on a spine not only a single but also two or three SBs can be found. We have rephrased this sentence for more clarity.

Page 9 Line 206 "L1 AZs did not show a large variability in size as indicated by the low SD, CV, and variance (Table 1)" Is this inter-patient variance of mean values? As in Supplementary Figure 1, both the SBs volume and PreAZ area show large variability in a given patient sample. Only the inter-patient variability of mean values seems low. Please state it clearly throughout the MS for other datasets as well.

For clarity concerning the variability between patients and structural parameters we have generated box plots (Suppl. Figures 1 and 2).

Page 9 Line 208 data is on Figure 5A and not 8A.

We thank the reviewer and corrected the citation of the Figure

Page 12 Line 295 how can the number of docked vesicles for L1b be larger than the one measured by the perimeter p10 nm? This later should contain the docked and PreAZ membrane proximal pool as well. This difference is even larger if we assume, that at EM tomography only partial AZs were analysed in a 200 nm thick section, not the entire AZ as for the perimeter measurement. Can the authors provide density estimates by dividing the docked / p10 nm vesicle numbers with the AZ area and comparing them?

This is a result comparing both methods. To the second concern: As stated in the text only synaptic boutons were the active zone can be followed from the beginning to its end and were the synaptic cleft was cut perpendicular were included in the TEM tomography sample as we also did in our 3D-volume reconstructions.

Methods:

Page 25 Line 624 While the PSD area can be equivocally measured, due to the dense appearance of the PSD on the EM images, the PreAZ is more difficult to outline due to lack of evident anatomical markers except the synaptic cleft (the dense material is much thinner). That is why in many publications the PreAZ area is considered to be identical to the PSD area. What are the anatomical criteria used here for the PreAZ? Why do the authors correct the PSD area, which is easy to measure with the PreAZ area that is much less certain to outline?

As stated in material and Methods both the pre- and postsynaptic densities are not defined by placing a closed contour in both densities because one can’t be certain that the dense accumulation of particles defining both areas since the impregnation (staining) and contrast of both structures critically depends on the uranyl and lead staining which could led to misinterpretation due to different staining results. That’s why we have drawn a contour line from the beginning to the end of the presynaptic density and extrapolated that for the postsynaptic density (for details see Material and Methods). In our samples both the pre- and postsynaptic densities were always clearly visible in those boutons further analyze.

Page 26 Line 640 vesicle density measurement: All the synaptic vesicles that are in the 50 nm thick section in their entirety are missed, and there are methods based on EM tomography to correct these estimations. One can not assume, that the error caused by "double counts" of vesicles cancels for the lost ones. There are stereological methods to estimate both types of error please include them and correct the values.

We would like to point out that the whole body of our work to structural analysis of vesicle pools is based on image data stemming from transmission electron microscopy (TEM) generating a projection of the entire volume of the ultra-thin section and NOT from scanning electron microscopy (SEM) where only a small volume close to the surface of the section would be captured. Operating in TEM mode ensures that no vesicle is missed only because it is embedded in its entirety in the section as postulated by the reviewer. Hence, EM tomography, which is basically a TEM operating from different incident angles in relation to the specimen or section, does not provide any advantage in detecting these vesicles. It does, however, help to better position a 3D object within the section volume itself and therefore allows to detect objects that could overlap from one viewing angle by using another angle. As the average vesicle diameter is of similar size compared to the section thickness, the possibility of a complete overlap to happen, however, is almost zero. And as we only count clear ring-like structures, a stereological correction factor calculated according to Abercrombie (1946) would underestimate real counts (see also Saetzler et al. 2002). If there is, however, relevant literature on "methods based on EM tomography" and "stereological methods to estimate both types of error" (over- and underestimates) that we are missing out on, we would appreciate the reviewer providing us with the corresponding references so that we can include such calculations in our paper.

Page 27 Line 664 and 665 "sections" are still tissue blocks, as sectioning comes after if the process is correctly written. Please correct.

We have corrected this according to the reviewer’s comment.

Page 43 Figure 4 D Data for L1b is missing, only the correlation line is visible.

Corrected in a new Figure.

Page 44 Figure 5 C arrowheads are in the correct places? Some of them do not seem to point to the edge of the synapse.

We carefully checked the Figure and adjusted the arrowheads.

Figure 5 E lower arrowhead labels something, that is difficult to identify but does not seem to be a vesicle.

We agree with the reviewer on this point and changed the figure accordingly.

Figure 5 F, the upper vesicle is at least 10 nm apart from the PreAZ membrane. Did the authors consider it as docked (indicated with arrowhead, according to the legend it labels docked vesicles)?

We agree with the reviewer on this point and changed the figure accordingly.

Page 45 Figure 6 B one of the 2 synaptic boutons (sb), sb2 has a tangential active zone that precludes the identification of the pre- and post-synaptic membranes, still 2 "docked vesicles" are labeled. How were they classified as docked? Please remove these tangential synapses from the dataset, as membranes can not be identified.

The reviewer is right that the active zone is tangentially cut, however, the two vesicles are associated with the AZ. In addition, we did not use this AZ for vesicle data analysis.

Page 46 Line 1124 interneuron axon labelled in green not brown.

Corrected as suggested by the reviewer.

Line 1129 SStC is missing.

Changed according to the reviewer’s comment.

Page 48 Table 2 Number of docked vesicles Median values are rounded to integer values? If yes why?

The statistic package used rounded to the given values.

Page 51 Supplementary Table 1 Hu_04 Histopathology, what does TLE stands for?

TLE: temporal lobe epilepsy. We included the abbreviation in the legend of Supplementary Table1, that is now table 2.

Reviewer #3 (Recommendations for the authors):

(1) Reanalysis of astrocytic coverage based on the % of synapses that form tripartite synapses.

We have reanalyzed the data concerning this point (new Figure 6D).

(2) Segregation of excitatory and inhibitory synapses.

We have now included a paragraph in our results section to distinguish between excitatory and inhibitory synapses.

(3) Better explanation of the limits of the study to assess functional parameters.

We disagree with the reviewer on this point and have not included an explanation concerning the limits of this study.

Changes from the earlier version to the revised version

Abstract

Abstract

Earlier version: Layer (L)1, beside receiving massive cortico-cortical, commissural and associational projections, is the termination zone of tufted dendrites of pyramidal neurons and the area of Ca2+ spike initiation. However, its synaptic organization in humans is not known. Quantitative 3D-models of synaptic boutons (SBs) in L1 of the human temporal lobe neocortex were generated from non-epileptic neocortical biopsy tissue using transmission electron microscopy, 3D-volume reconstructions and EM tomography. Particularly, the size of active zones (AZs) and the readily releasable, recycling and resting pool of synaptic vesicles (SVs) were quantified.

Revised version: Layer (L)1, beside receiving massive cortico-cortical, commissural and associational projections, is the termination zone of tufted dendrites of pyramidal neurons and the area of Ca2+ spike initiation. However, its synaptic organization in humans is not known. Quantitative 3D-models of synaptic boutons (SBs) in L1 of the human temporal lobe neocortex were generated from non-epileptic neocortical biopsy tissue using transmission electron microscopy, 3D-volume reconstructions and TEM tomography. Particularly, the size of active zones (AZs) and the readily releasable, recycling and resting pool of synaptic vesicles (SVs) were quantified.

Earlier version: SBs had a single large AZ (∼0.20 µm2), a total pool of ∼3500 SVs, a large readily releasable (∼4 SVs), recycling (∼470 SVs) and resting (∼2900 SVs) pool. Astrocytic coverage suggests cross talk at synaptic complexes.

Revised version: The majority of SBs contained numerous mitochondria comprising ∼7% of the total volume of SBs, had a large macular, non-perforated AZs (∼0.20 µm2) and were predominantly located on dendritic spines. SBs had a total pool of ∼3500 SVs, a comparably large readily releasable (∼4 SVs), recycling (∼470 SVs) and resting (∼2900 SVs) pool. Astrocytic coverage of SBs suggests both synaptic cross talk or removal of spilled glutamate by astrocytic processes at synaptic complexes.

Earlier version: Thus, L1 SBs mediate, integrate and synchronize contextual and cross-modal information, enabling flexible and state-dependent processing of feedforward sensory inputs from other layers of the cortical column.

Revised version: The structural composition of SBs in L1 may underlie the function of L1 networks that mediate, integrate and synchronize contextual and cross-modal information, enabling flexible and state-dependent processing of feedforward sensory inputs from other layers of the cortical column.

Introduction

The mammalian neocortex is the most complex part of the brain comprising over 75% of the gray matter. One fundamental feature established during development is the formation of distinct layers with different networks of neurons, and its organization into vertically oriented slabs, so-called cortical columns (Marin-Padilla 1978, 1998; reviewed by Rockland and DeFelipe 2018). This organization is established in an ‘inside layers first-outside layers last’ fashion (Luskin and Shatz 1985a, b; reviewed by Cooper et al. 2008). The exception, beside layer (L)6 is L1 generated first originating from the marginal zone (Marin Padilla 1998; reviewed by Bystron et al. 2008). Ongoing cortico- and synaptogenesis finally lead to the formation of the adult cortical network with its unique structural and functional properties (reviewed by Lübke and Feldmeyer 2007; Rockland and DeFelipe 2018).

Earlier version: Remarkably, L1 is highly conserved across cortical areas and higher mammalian species, but its importance was long underestimated. L1 is considered as the predominant input layer for top-down information, relayed by a rich, dense network of cortico-cortical, cortico-thalamic, commissural and associational long-range axonal projections. In humans, and thus different to experimental animals, both L1 excitatory Cajal-Retzius and a hetregeneous population of inhibitory neurons provide signals to the terminal tuft dendrites of pyramidal neurons of the underlying cortical layers. Thus, L1 is a central locus of neocortical associations controlled by distinct types of inhibition and feedforward excitation (Schumann et al. 2019; Hartung and Letzkus 2021, reviewed by Schumann et al. 2021).

Revised version: Remarkably, L1 is highly conserved across cortical areas and higher mammalian species, but its importance was long underestimated. However, L1 is now considered as the predominant input layer for top-down information, relayed by a rich, dense network of cortico-cortical, cortico-thalamic, commissural and associational long-range axonal projections. In humans, and thus different to experimental animals, both L1 excitatory Cajal-Retzius (CR) cells (Anstötz et al. 2014) and a quite heterogeneous population of inhibitory neurons (see for example Jiang et al. 2013; Boldog et al. 2018; Obermayer et al. 2018; Kwon et al. 2019; reviewed by Huang et al. 2024) providing signals to the terminal tuft dendrites of pyramidal neurons of the underlying cortical layers. Thus, L1 is a central locus of neocortical associations controlled by distinct types of inhibition and feedforward excitation (Schumann et al. 2019; Hartung and Letzkus 2021, reviewed by Schumann et al. 2021).

Earlier version: The temporal lobe neocortex (TLN) is situated on the basolateral aspect of the cerebral hemispheres primarily observed in primates, including humans. Notably, in human, the TLN encompasses approximately 20% of the entire cerebral cortex (Kiernan 2012). It is recognized as a highly specialized associative neocortex, characterized by its homotypic granular six-layered organization (von Economo and Koskinas 1925; Palomero-Gallagher and Zilles 2019). The human TLN plays a crucial role in various cognitive functions, including auditory, visual, vestibular, linguistic, and olfactory processing. Additionally, it is intricately connected to diverse sensory and multimodal associational brain regions, such as the limbic system, amygdala, and various subcortical structures (Insausti 2013). Lastly, the human TLN serves as the epicenter and onset site for temporal lobe epilepsy (Allone et al. 2017; Tai et al. 2018).

Revised version: The temporal lobe neocortex (TLN) is situated on the basolateral aspect of the cerebral hemispheres primarily observed in higher primates, including humans. Notably, in human, the TLN encompasses approximately 20% of the entire cerebral cortex (Kiernan 2012). It is recognized as a highly specialized associative neocortex, characterized by its homotypic granular six-layered organization (von Economo and Koskinas 1925; Palomero-Gallagher and Zilles 2019). The human TLN plays a crucial role in various cognitive functions, including auditory, visual, vestibular, linguistic, and olfactory processing. Additionally, it is intricately connected to diverse sensory and multimodal associational brain regions, such as the limbic system, amygdala, and various subcortical structures (Insausti 2013). Lastly, the human TLN serves as the epicenter and onset site for temporal lobe epilepsy (Allone et al. 2017; Tai et al. 2018).

Earlier version: One of the most important discoveries in neuroscience, beside the definition of the neuron by Cajal’s neuronal doctrine (1911), was the introduction of the term ‘Synapse’, more than 100 years ago. Since then these structures have been investigated from different viewpoints including structural, functional, molecular and computational studies summarized in meanwhile thousands of original publications, reviews and numerous textbooks. Yet, the quantitative geometry of the most prevalent synapse in the brain, the cortical synapse, has remained largely unexplored structurally, in both experimental animal models (but see Rollenhagen et al. 2015, 2018; Bopp et al. 2017; Hsu et al. 2017; Prume et al. 2020). and, even more so, in the context of the human brain (but see Yakoubi et al. 2019a, b; Dominguez-Alvaro et al. 2019, 2021; Cano-Astorga et al. 2021, 2023; Schmuhl-Giesen et al. 2022). This could be partially attributed to the availability of human brain tissue samples. Nevertheless, detailed and quantitative descriptions of synaptic complexes in L1 of the human brain are to date non-existing.

Revised version: One of the most important discoveries in neuroscience, beside the definition of the neuron by Cajal’s neuronal doctrine (1911), was the introduction of the term ‘Synapse’, more than 100 years ago. Since then these structures have been investigated from different viewpoints including structural, functional, molecular and computational studies summarized in meanwhile thousands of original publications, reviews and numerous textbooks. Yet, the quantitative geometry of the most prevalent synapse in the brain, the cortical synapse, has remained largely unexplored structurally, in both experimental animals (but see Rollenhagen et al. 2015, 2018; Bopp et al. 2017; Hsu et al. 2017; Prume et al. 2020) and, even more so, in the context of the human brain (but see Yakoubi et al. 2019a, b; Dominguez-Alvaro et al. 2019, 2021; Cano-Astorga et al. 2021, 2023; Schmuhl-Giesen et al. 2022; Shapson-Coe et al. 2024). This could be partially attributed to the availability of human brain tissue samples. Nevertheless, detailed and quantitative descriptions of synaptic complexes in L1 of the human brain are to date non-existing.

Earlier version: Here, an investigation into the synaptic organization of L1 in the human TLN was undertaken using high-resolution transmission electron microscopy (TEM), 3D-volume reconstruction to generate quantitative 3D-models of synaptic boutons (SBs) and by electron microscopic (EM) tomography. The ultimate objective of this comprehensive project is to elucidate the synaptic organization, layer by layer, of the cortical column in humans.

Revised version: Here, an investigation into the synaptic organization of L1 in the human TLN was undertaken using high-resolution transmission electron microscopy (TEM), 3D-volume reconstruction to generate quantitative 3D-models of synaptic boutons (SBs) and by TEM tomography. The ultimate objective of this comprehensive project is to elucidate the synaptic organization, layer by layer, of the cortical column in humans.

Such meticulous quantitative structural analyses of SBs and their corresponding target structures are indispensable for comprehending, elucidating, and establishing connections between the structural and functional properties of the diverse signal cascades underlying synaptic transmission, efficacy, strength, and plasticity. Consequently, these efforts contribute significantly to understanding the computational properties of brain networks, with the organization of the cortical column emerging as a pivotal component thereof.

Results

Added in the revised version: Neural and synaptic composition of L1 in the human TLN

Added in the revised version: L1 in the adult mammalian neocortex represent a relatively cell sparse layer that in humans can be subdivided into two sublaminae, L1a and L1b as revealed in semithin sections (Figure 1A) and Golgi-preparations (Figure 1B, C, F, G) and at the EM level (Figure 1H, I). L1a is more dominated by a dense network of astrocytes of different shape and size, their fine processes (Figure 2A), the occurrence of reactive microglia (Figures 1H, I, 2A) and terminal tuft dendrites of pyramidal neurons throughout the neuropil of L1 (Figures 1H, I, 2B).

Added in the revised version: L1 and L2 in the human TLN as revealed in methylene-blue stained semithin sections, Golgi-impregnation and low power EM.

Added in the revised version: A, Methylene-blue stained semithin section of L1a, L1b, and L2. The cell sparse zone contains degenerating neurons (dark appearance, arrowheads in the framed area) and a cell body (asterisk in the framed area). Note the sudden increase in the density of neurons at the L1b/L2 border. Scale bar 500 µm. B, C, Golgi-impregnated radial astrocytes in L1a (B) and L1b (C). Scale bars 50 µm. D, Group of neurons in L1b containing lipofuscin granules of various shape and size (asterisks) most prominent in two of the neurons. Scale bar 50 µm. E, Putative CR-cell identifiable by its long horizontally oriented dendrite (marked by arrowheads) and an inverted putative GABAergic interneuron (asterisk) in L1b. Scale bar 50 µm. F, G, Golgi-impregnated CR-Cells in L1a with the characteristic horizontal orientation of dendrites (arrowheads) and axons (asterisks). Scale bars 50 µm. H, EM micrograph of L1a underneath the pial surface containing two putative CR-cells (asterisks). Note the abrupt transition (indicated by the dashed line) to the neuropil of L1b, and a reactive microglial cell (arrowhead), and dendrites and synapses. Scale bar 25 µm. I, EM micrograph of L1b containing thousands of dendritic profiles, a putative GABAergic interneuron (asterisk) but also degenerating neurons (arrowheads) and a reactive microglia (I1) Scale bar 50 µm.

Added in the revised version: Structural and synaptic organization of L1 in the human TLN as revealed by EM.

Added in the revised version: A, Superficial part of L1a composed of a dense network of fine astrocytic processes (asterisks). In the reactive microglial cell (rmic) so-called lipofuscin granules (framed area) were observed. Below the astrocytic network synaptic complexes gradually increased (highlighted in blue for target structures, yellow for SBs). Scale bar 10 µm. B, Two apical dendrites (blue) traversing L2 with beginning terminal tuft dendrites in L2 and L1 contacted by SBs (transparent yellow); some of which are putatively GABAergic terminating on the dendritic shaft (asterisks). Scale bar 2.5 µm. C, Large GABAergic SB (sb, transparent yellow) identifiable by the small-sized ovoid SVs and the less prominent PSD (arrowheads) synapsing on a dendritic shaft (de, transparent blue) in L1b. D, Three SBs in L1a (sb1-sb3, transparent yellow) located on a dendritic spine (sp, transparent blue) and a neighboring dendritic shaft (sh, transparent blue). Note that sb1 is excitatory as indicated by the larger round SVs with a prominent AZ whereas sb2 and sb3 are putative inhibitory terminals identified by smaller ovoid SVs and the lack of a prominent PSD. E, Elongated spine (sp) emerging from a small caliber dendrite (de, transparent blue) receiving input from a small-sized SB (sb, transparent yellow). Note the large, elongated spine apparatus (framed area). Scale bar in C-E 0.5 µm. In graphs C-E AZs are marked by arrowheads.

Added in the revised version: Hence, L1a can be described as a sublamina with a predominant proportion of astrocytic and microglial processes, although it also harbors numerous other structural elements such as terminal tuft dendrites and synaptic complexes. These complexes consist of a SB juxtaposed with either a dendritic shaft or spine (Figure 2A) directly beneath or interwoven among the fine astrocytic processes within L1 (Figure 2A). Lipofuscin granules, indicative of aging, were frequently noted in astrocytes, reactive microglia (Figure 2A) and L1 neurons, although their abundance, size, and morphology varied considerably.

Added in the revised version: Remarkably, and in contrast to rodents a persistent subpopulation of excitatory CR-cells were found in humans identifiable by the horizontal bipolar orientation of both dendrites and axons (Figure 1E-H). The majority were located directly underneath the pial surface (Figure 1F, G) or intermingled with a heterogeneous population of GABAergic interneurons (Figure 1E).

Added in the revised version: In the human TLN, GABAergic interneurons constitute the predominant class of neurons in L1, exhibiting significant structural and functional heterogeneity (Figure 1D, E; see also Jiang et al, 2014, Verhoog et al. 2016; Boldog et al. 2018; Obermayer et al. 2018; Kwon et al. 2019; Schuman et al. 2019, reviewed by Huang et al. 2024) which were found throughout both sublaminae and were sometimes organized in clusters of 3-6 neurons (Figure 1D). Beside CR-cells and GABAergic interneurons, also degenerating neurons, identifiable by their shrunken and dark appearance (Figure 1A, I) were found in both sublaminae although their numbers varied substantially between brain tissue samples.

Added in the revised version: L1b is primarily dominated by dendritic and synaptic profiles. due to a massive increase in synaptic complexes (Figures 1H, 2B). Furthermore, L1b is infiltrated by apical dendrites of L2, L3 and L5 pyramidal neurons which are interspersed with smaller dendritic segments representing apical oblique dendrites (Figures 1I, 2B). The majority of synaptic complexes in L1 were axo-spinous (∼80%, Figure 2 B, D, E), the remainder were axo-dendritic some of which but not all were regarded as putative GABAergic terminals (Figure 2C; but see Silver et al. 2003).

Added in the revised version: L1b can be separated from L2 by the sudden increase of neuronal cell bodies and thick apical dendrites originating at the upper pole of the somata of pyramidal neurons (Figures 1A, 2B). At the L1b and L2 transition more and more apical and apical oblique dendrites of different caliber were observable.

Added in the revised version: Synaptic organization of L1 in the human TLN

Added in the revised version: The primary objective of this study was to investigate the synaptic organization of excitatory SBs in L1 in the human TLN, focusing on synaptic parameters that represent possible morphological correlates of synaptic transmission, efficacy, strength and plasticity. The quantitative analysis was separated for both sublaminae to look for possible sublaminar-specific differences. Beside the overall geometry of SBs and the number and size of mitochondria, in particular the size of the PreAZ and PSD constituting the AZ, the morphological equivalent to a functional neurotransmitter release site and the organization of the three functionally defined pools of SVs, namely the readily releasable (RRP), the recycling (RP) and resting pool, were quantified.

Added in the revised version: To achieve this, a total of 361 SBs was completely reconstructed out of four patients in six series containing 100-150 ultrathin sections/series. A total of 190 SBs and 184 AZs (L1a) and 171 SBs and 165 AZs (L1b) using non-affected neocortical access tissue (see Material and Methods; Source Data 02) were completely 3D-reconstructed and quantified using fine-scale TEM, digital TEM images and subsequent 3D-volume reconstructions.

Added in the revised version: Excitatory and inhibitory synaptic complexes in L1a und L1b were formed by either presynaptic en passant or endterminal boutons with their prospective postsynaptic target structures dendritic shafts (Figures 2B C, D, 3A-C) or spines of different caliber and type (Figures 2B, D, E, 3A-C). In both sublaminae, SBs were predominantly found on dendritic spines (∼81%; L1a) and ∼79% (L1b), the remainder were located on dendritic shafts. The majority of SBs on spines were located on mushroom spines (L1a: ∼83%, L1b: ∼73%), a smaller fraction on stubby spines (L1a: ∼11%, L1b: ∼9%) or on thin elongated spines;(L1a: 6%, L1b: ∼22%; Figure 2E), the remainder were not classifiable.

Added in the revised version: Since our focus was on excitatory SBs, inhibitory terminals were not further analyzed in detail. Inhibitory SBs were distinguished from excitatory SBs by their smaller SVs which also has a more spherical appearance and the smaller sized PreAZ and the lack of a prominent PSD (Figure 2D). Infrequently, GABAergic and glutamatergic terminals were found on the same dendrite (Figure 2C, D) or spine (see for example Wittner et al. 2001) but rare by ∼2-3%. (Kwon et al. 2019). In turn, also not all SBs located on dendritic shafts were GABAergic (see for example Silver et al. 2003).

Added in the revised version: The ratio between excitatory vs. inhibitory SBs was between 10-15% although with layer-specific differences. Our findings are in good agreement with recent publications by Shapson-Coe et al. (2024) and by Cano-Astorga et al. (2023, 2024) which also estimated the number of inhibitory SBs in different layers of the temporal lobe neocortex by10-15%. However, no significant differences were found in the content of inhibitory SBs between the two sublaminae L1a and L1b.

Added in the revised version: Some SBs in both sublaminae were seen to establish either two or three SBs on the same spine, spines of other origin or dendritic shafts. Remarkably, ∼90% of spines in L1a and L1b contained a spine apparatus (Figure 2E), a specialized form of the endoplasmic reticulum, structures that increase spine motility and may also stabilize the synaptic complex during signal transduction (Deller et al. 2003; reviewed by Knott and Holtmaat 2008).

Added in the revised version: Furthermore, in numerous SBs, so-called dense-core vesicles (DCVs) were distributed throughout the presynaptic terminal intermingled with the population of SVs (Figure 3B, C, E).

Added in the revised version: 3D-volume reconstructions of SBs and vesicle pools in L1a and L1b of the human TLN.

Added in the revised version: A-D, Here, all SBs are colored in transparent yellow. Subelements: PreAZs (red), SVs (green), DCVs (magenta), mitochondria (white). The postsynaptic target structures are given in blue A, Six SBs (sb1-sb6) terminating at different locations on a dendritic segment in L1b, two of which (syn3 and syn4) were located on spines, the remaining on the dendritic shaft. Scale bar 1 μm. B, Large SB terminating on two dendritic segments of different shape and size (de1, de2) in L1a. C, SB synapsing on a dendrite (de) and spine (sp) in L1b containing several mitochondria associated with the pool of SVs. D, E, 3D-volume reconstructions of the total pools of SVs (green dots). Note the different shape and size of the PreAZs (red) with either a non-perforated macular (D) or perforated (E) appearance and the comparably large total pool size. Note the large DCVs intermingled with the pool of SVs.

Added in the revised version: Remarkably, beside synaptic complexes, also ‘special connections’ in L1 of the human TLN were observed. This encompassed so-called dendro-dendritic synapses and fine astrocytic processes that receive synaptic input from SBs and in turn, provide synaptic output to either dendritic shafts or spines (not shown, but see Yakoubi et al. 2019a, b; Schmuhl-Giesen et al. 2022).

Synaptic density in L1a and L1b of the human TLN

Synaptic density measurements serve as a valuable tool for quantitatively characterizing the synaptic organization of a specific cortical layer, assessing the connectivity rate, and identifying potential inter-individual differences among patients within the human TLN.

Earlier version: The average density of synaptic complexes in L1 was comparatively high with a value of 5.26*108 ± 8.44*107, 5.52*108 ± 1.86*107 synapses/mm3 in L1a and 3.93*108 ± 1.35*107 synapses/mm3 in L1b, ranging from 3.9*108 to 7.3*107 (L1a) and from 2.6*108 to 5.3*108 (L1b), respectively (Source Data 1). Interestingly, the average synaptic density in L1a was by ∼1.5-fold larger when compared with L1b although L1a was regarded as the more astrocytic dominated sublamina. Strikingly, a huge inter-individual variability was found, although with no significant difference between L1a and L1b.

Revised version: The average density of synaptic complexes in L1 was comparatively high with a value of 5.26*108 ± 8.44*107, 5.52*108 ± 1.25*108 synapses/mm3 in L1a and 5.01*108 ± 1.40*107 synapses/mm3 in L1b, ranging from 4.08*108 to 6.31*108 (L1a) and from 4.87*108 to 5.15*108 (L1b), respectively (Supplemental Table 1; Source Data 01). Strikingly, a huge inter-individual variability was found, although with no significant difference between L1a and L1b.

Earlier version: Therefore, L1 exhibited a high density of synaptic complexes, indicating a relatively robust connectivity of neurons within L1 of the human TLN.

Revised version: Therefore, L1 exhibited a comparably high density of synaptic complexes, indicating a relatively robust and thus strong connectivity of neurons within L1 of the human TLN.

Removed in the revised version: Neural and synaptic composition of L1 in the human TLN

Removed in the revised version: L1 in the adult mammalian neocortex represent a relatively cell sparse layer that in humans can be subdivided into two sublaminae, L1a and L1b as revealed in Golgi-preparations, semithin sections (Figure 1A) and at the EM level (Figure 1H, I). L1a is dominated by a dense network of astrocytes of different shape and size, their fine processes (Figure 2A), the occurrence of reactive microglia (Figures 1H, I, 2A) and terminal tuft dendrites of pyramidal neurons throughout the neuropil of L1 (Figures 1H, I, 2B).

Removed in the revised version: L1 and L2 in the human TLN as revealed in methylene-blue stained semithin sections, Golgi-impregnation and low power EM.

Removed in the revised version: A, Methylene-blue stained semithin section of L1a, L1b, and L2. The cell sparse zone contains degenerating neurons (dark appearance, arrowheads in the framed area) and a cell body (asterisk in the framed area). Note the sudden increase in the density of neurons at the L1b/L2 border. Scale bar 500 µm. B, C, Golgi-impregnated radial astrocytes in L1a (B) and L1b (C). Scale bars 50 µm. D, Group of neurons in L1b containing lipofuscin granules of various shape and size (asterisks) most prominent in two of the neurons. Scale bar 50 µm. E, Putative CR-cell identifiable by its long horizontally oriented dendrite (marked by asterisks) and an inverted putative GABAergic interneuron (asterisk) in L1b. Scale bar 50 µm. F, G, Golgi-impregnated CR-Cells in L1a with the characteristic horizontal orientation of dendrites (arrowheads) and axons (asterisks). Scale bars 50 µm. H, EM micrograph of L1a (astrocytic layer, astl) underneath the pial surface. Note the abrupt transition (indicated by the dashed line) to a neuropil (neurl) containing two putative CR-cells (asterisks), a reactive microglial cell (arrowhead), and dendrites and synapses. Scale bar 25 µm. I, EM micrograph of L1b containing thousands of dendritic profiles, a putative GABAergic interneuron (framed area, asterisk) but also degenerating neurons (framed area, arrowheads) and a reactive microglia (I1) Scale bar 50 µm.

Removed in the revised version: Structural and synaptic organization of L1 in the human TLN as revealed by EM.

Removed in the revised version: A, Superficial part of L1a composed of a dense network of fine astrocytic processes (asterisks). In the reactive microglial cell (rmic) so-called lipofuscin granules (framed area) were observed. Below the astrocytic network synaptic complexes gradually increased (highlighted in blue for target structures, yellow for SBs). Scale bar 10 µm. B, Two apical dendrites (blue) traversing L2 with beginning terminal tuft dendrites in L2 and L1 contacted by SBs (transparent yellow); some of which are putatively GABAergic terminating on the dendritic shaft (asterisks). Scale bar 2.5 µm. C, Large GABAergic SB (sb, transparent yellow) identifiable by the small-sized ovoid SVs and the less prominent PSD (arrowheads) synapsing on a dendritic shaft (de, transparent blue) in L1b. D, Three SBs in L1a (sb1-sb3, transparent yellow) located on a dendritic spine (sp, transparent blue) and a neighboring dendritic shaft (sh, transparent blue). Note that sb1 is excitatory as indicated by the larger round SVs with a prominent AZ whereas sb2 and sb3 are putative inhibitory terminals identified by smaller ovoid SVs and the lack of a prominent PSD. E, Elongated spine (sp) emerging from a small caliber dendrite (de, transparent blue) receiving input from a small-sized SB (sb, transparent yellow). Note the large, elongated spine apparatus (framed area). Scale bar in C-E 0.5 µm. In graphs C-E AZs are marked by arrowheads.

Removed in the revised version: Hence, L1a can be characterized as a predominantly astrocytic sublamina, whereas L1b is primarily dominated by dendritic and synaptic profiles. Lipofuscin granules, indicative of aging, were frequently noted in astrocytes, reactive microglia (Figure 2A) and neurons, although their abundance, size, and morphology varied considerably.

Removed in the revised version: Remarkably, and in contrast to rodents in humans a persistent subpopulation of CR-cells were found identifiable by the horizontal bipolar orientation of both dendrites and axons (Figure 1E-H). The majority were found directly underneath the pial surface (Figure 1F, G) or intermingled with a heterogeneous population of GABAergic interneurons (Figure 1E).

Removed in the revised version: In the human TLN, GABAergic interneurons constitute the predominant class of neurons in L1, exhibiting significant structural heterogeneity, but structurally very heterogeneous class of neurons in L1 (Figure 1D, E; see also Verhoog et al. 2016; Schuman et al. 2019, reviewed by Schuman et al. 2021) which were found throughout both sublaminae and were sometimes organized in clusters of 3-6 neurons (Figure 1D).

Removed in the revised version: Beside CR-cells and GABAergic interneurons, also degenerating neurons, identifiable by their shrunken and dark appearance (Figure 1A, I) were found in both sublaminae although their numbers varied substantially between brain tissue samples.

Removed in the revised version: While L1a is primarily characterized by the presence of astrocytes and their fine processes, it also harbors numerous synaptic complexes. These complexes consist of a SB juxtaposed with either a dendritic shaft or spine (Figure 2A) directly beneath or interwoven among the fine astrocytic processes within L1 (Figure 2A).

Removed in the revised version: In contrast, L1b is a more ‘dendritic and synaptic layer’ due to a massive increase in synaptic complexes (Figures 1H, 2B). Furthermore, L1b is infiltrated by apical dendrites of L2, L3 and L5 pyramidal neurons which are interspersed with smaller dendritic segments representing apical oblique dendrites (Figures 1I, 2B). The majority of synaptic complexes were axo-spinous (∼80%, Figure 2 B, D, E), the remainder were axo-dendritic some of which but not all were regarded as putative GABAergic terminals (Figure 2C).

Removed in the revised version: L1b can be separated from L2 by the sudden increase of neuronal cell bodies and thick apical dendrites originating at the upper pole of the somata of pyramidal neurons (Figures 1A, 2B). At the L1b and L2 transition more and more apical and apical oblique dendrites of different caliber were observable.

Removed in the revised version: Synaptic organization of L1 in the human TLN

Removed in the revised version: The primary objective of this study was to explore the synaptic organization of L1 in the human TLN, focusing on synaptic parameters that represent morphological indicators of synaptic transmission, efficacy, strength and plasticity. The quantitative analysis was conducted separately for both sublaminae to identify potential sublaminar-specific differences. The quantitative analysis was separated for both sublaminae to look for possible sublaminar-specific differences. Beside the overall geometry of SBs and the number and size of mitochondria, in particular the size of the PreAZ and PSD constituting the AZ, the morphological equivalent to a functional neurotransmitter release site and the organization of the three pools of SVs, namely the readily releasable (RRP), the recycling (RP) and resting pool, were quantified.

Removed in the revised version: To achieve this, a total of 361 SBs was completely reconstructed out of four patients in six series containing 100-150 ultrathin sections/series. A total of 190 SBs and 184 AZs (L1a) and 171 SBs and 165 AZs (L1b) using non-affected neocortical access tissue (see Material and Methods; Source Data 2) were completely 3D-reconstructed and quantified using file-scale TEM and subsequent 3D-volume reconstructions.

Removed in the revised version: Excitatory and inhibitory synaptic complexes in L1a und L1b were formed by either presynaptic en passant or endterminal boutons with their prospective postsynaptic target structures dendritic shafts (Figures 2B C, D, 3A-C) or spines of different caliber and type (Figures 2B, D, E, 3A-C). In both sublaminae, SBs were predominantly found on dendritic spines (∼81%; L1a) and ∼79% (L1b), the remainder were located on dendritic shafts. SBs on spines were predominantly located on mushroom spines (L1a: ∼83%, L1b: ∼73%), a smaller fraction on stubby spines (L1a: ∼11%, L1b: ∼9%) or on thin elongated spines;(L1a: 6%, L1b: ∼22%; Figure 2E), the remainder were not classifiable.

Removed in the revised version: Numerous SBs in both sublaminae were seen to establish either two or three synaptic contacts on the same spine or dendrite. Infrequently, GABAergic synapses identified by the smaller more spherical SVs and glutamatergic terminals were found on the same dendrite (Figure 2C, D) or spine. Remarkably, ∼90% of spines in L1a and L1b contained a spine apparatus (Figure 2E), a specialized form of the endoplasmic reticulum, structures that increase spine motility and may also stabilize the synaptic complex during signal transduction (Deller et al. 2003; reviewed by Knott and Holtmaat 2008).

Removed in the revised version: In numerous SBs, so-called dense-core vesicles (DCWs) were distributed throughout the presynaptic terminal intermingled with the population of SVs (Figure 3B, C, E).

Removed in the revised version: 3D-volume reconstructions of SBs and vesicle pools in L1a and L1b of the human TLN. A-D,

Removed in the revised version: Here, all SBs are colored in transparent yellow. Subelements: PreAZs (red), SVs (green), DCVs (magenta), mitochondria (white). The postsynaptic target structures are given in blue A, Six SBs (sb1-sb6) terminating at different locations on a dendritic segment in L1b, two of which (syn3 and syn4) were located on spines, the remaining on the dendritic shaft. Scale bar 1 μm. B, Large SB terminating on two dendritic segments of different shape and size (de1, de2) in L1a. C, SB synapsing on a dendrite (de) and spine (sp) in L1b containing several mitochondria associated with the pool of SVs. D, E, 3D-volume reconstructions of the total pools of SVs (green dots). Note the different shape and size of the PreAZs (red) with either a non-perforated macular (D) or perforated (E) appearance and the comparably large total pool size. Note the large DCVs intermingled with the pool of SVs.

Removed in the revised version: Remarkably, beside synaptic complexes, also ‘special connections’ in L1 of the human TLN were observed. This encompassed so-called dendro-dendritic synapses and fine astrocytic processes that receive synaptic input from SBs and in turn, provide synaptic output to either dendritic shafts or spines (not shown, but see Yakoubi et al. 2019a, b: Schmuhl-Giesen et al. 2022),

Geometry and size of SBs and mitochondria in the human TLN

Earlier version: Overall, L1 SBs were on average medium-sized, with a mean surface area of 5.48 ± 1.40 µm2, and a mean volume of 0.50 ± 0.19 µm3, with a slight difference in size between L1a and L1b. Interestingly, the variability in both surface area and volume of SBs was comparably small in L1a and L1b as indicated by a low CV and variance (Table 1; Source Data 2) regardless of their target structures. SBs in L1 were comparable in size with SBs in L5, L6, but ∼2-fold larger than those in L4 of the human TLN (p ≤ 0.001; see also Yakoubi et al. 2019a, b: Schmuhl-Giesen et al. 2022),)

Revised version: Overall, L1 SBs were on average medium-sized, with a mean surface area of 5.48 ± 1.40 µm2, and a mean volume of 0.50 ± 0.19 µm3, with a slight difference in size between L1a and L1b. Interestingly, the variability in both surface area and volume of SBs was comparably small in L1a and L1b as indicated by a low CV and variance (Table 1; Source Data 02) regardless of their target structures. SBs in L1 were comparable in size with SBs in L5, L6, but ∼2-fold larger than those in L4 of the human TLN (p ≤ 0.001; see also Yakoubi et al. 2019a, b: Schmuhl-Giesen et al. 2022).

Earlier version: Mitochondria play a pivotal role in synaptic transmission and plasticity (reviewed by Dallerac et al. 2018). In L1 SBs either no (Figures 2B, 3A, E) or several mitochondria (Figures 2, 3C, D; range 0 to 8) of different shape and size were observed. Mitochondria had a volume of 0.04 ± 0.02 µm3 in L1a and 0.05 ± 0.01 µm3 in L1b, respectively (Table 1; Source Data 2). Mitochondria contributed with ∼7% (L1) to the total bouton volume, similar to values in L6 (∼6%; Schmuhl-Giesen et al. 2022), but with ∼1.5-fold lower percentage as those in L5 (∼12%; Yakoubi et al. 2019a) and L4 (∼13%; Yakoubi et al. 2019b).

Revised version: Mitochondria play a pivotal role in synaptic transmission and plasticity (reviewed by Dallerac et al. 2018; see also Discussion). In L1 SBs either no (Figures 2B, 3A, E) or several mitochondria (Figures 2, 3C, D; range 0 to 8) of different shape and size were observed. When present mitochondria had a volume of 0.04 ± 0.02 µm3 in L1a and 0.05 ± 0.01 µm3 in L1b, respectively (Table 1; Source Data 02). Mitochondria contributed with ∼7% (L1) to the total bouton volume of SBs, similar to values in L6 (∼6%; Schmuhl-Giesen et al. 2022), but with a ∼1.5-fold lower percentage as those in L5 (∼12%; Yakoubi et al. 2019a) and L4 (∼13%; Yakoubi et al. 2019b).

Earlier version: A good correlation between the volume of the SBs and that of mitochondria was found for L1a (R2: 0.7528; Figure 4B) and for L1b (R2: 0.5992; Figure 4B). A weak correlation was also found for L6 (Schmuhl-Giesen et al. 2022) but a strong one in both L4 (Yakoubi et al. 2019a) and L5 (Yokoubi et al. 2019b) suggesting a layer-specific difference in the content of mitochondria in individual SBs.

Revised version: A good correlation between the volume of the SBs and that of mitochondria was found for L1a (R2: 0.7528; Figure 4B) and for L1b (R2: 0.5992; Figure 4B). A weak correlation was also found for L6 (Schmuhl-Giesen et al. 2022) but a strong one in both L4 (Yakoubi et al. 2019a) and L5 (Yakoubi et al. 2019b) suggesting a layer-specific difference in the content of mitochondria in individual SBs.

Correlation plots for structural and synaptic parameters characterizing L1 SBs.

Earlier version: Correlation plots showing the strength of correlations between structural and synaptic parameters (Source Data 2). Correlations dots and regression lines for L1a are given in black and that for L1b in red. A, Surface area of SBs vs. surface area of PreAZ; B, Volume of SBs vs. volume of mitochondria; C, Surface area of PreAZs vs. total pool of SVs; D, Surface area of PreAZs vs. p10 nm RRP; E, Surface area of PreAZs vs. p20 nm RRP; F, Surface area of PreAZs vs. p60-p200 nm RP; G, Surface area of PreAZs vs. p200 nm resting pool; H, Surface area of PreAZs vs. p500 nm resting pool.

Revised version: Correlation plots showing the strength of correlations between structural and synaptic parameters (Source Data 02). Correlations dots and regression lines for L1a are given in black and that for L1b in red. A, Surface area of SBs vs. surface area of PreAZs; B, Volume of SBs vs. volume of mitochondria; C, Surface area of PreAZs vs. total pool of SVs; D, Surface area of PreAZs vs. p10 nm RRP; E, Surface area of PreAZs vs. p20 nm RRP; F, Surface area of PreAZs vs. p60-p200 nm RP; G, Surface area of PreAZs vs. p200 nm resting pool; H, Surface area of PreAZs vs. p500 nm resting pool.

Structural composition of AZs in L1a and L1b excitatory SBs in the human TLN

Earlier version: The number, size, and shape of the AZ, composed of the PreAZ and PSD, is one key structural determinant in synaptic transmission and plasticity (Südhof 2002; Matz et al. 2010; Holderith et al. 2012) as it represents the docking area of SVs (PreAZs) and the receiving area of neurotransmitter quanta (PSD) (Südhof 2012). The majority (∼98%) of SBs in L1a and L1b had only a single (Figures 2C-E, 3A-C, E) at most two or three AZs (Figure 3D, D). Beside very large, spanning the entire pre- and postsynaptic apposition zone, also smaller AZs covering only a fraction of the pre- and postsynaptic apposition zone were found. The majority of AZs (∼75%) were of the macular, non-perforated type (Figures 2C-E, 3A-D), the remainder showed either a perforation in the PreAZ, PSD or both (Figure 3E).

Revised version: The number, size, and shape of the AZ, composed of the PreAZ and PSD, is one key structural determinant in synaptic transmission and plasticity (Südhof 2002; Matz et al. 2010; Holderith et al. 2012) as it represents the docking area of SVs (PreAZs) and the receiving area of neurotransmitter quanta at the PSD (Südhof 2012).

Added in the revised version: The majority (∼98%) of SBs in L1a and L1b had only a single (Figures 2C-E, 3A-C, E) at most two or three AZs (Figure 3D). Beside very large, spanning the entire pre- and postsynaptic apposition zone, also smaller AZs covering only a fraction of the pre- and postsynaptic apposition zone were found. The majority of AZs (∼75%) were of the macular, non-perforated type (Figures 2C-E, 3A-D), the remainder showed either a perforation in the PreAZ, PSD or both (Figure 3E).

Earlier version: On average, PreAZs were 0.18 ± 0.08 µm2 in surface area in L1a and 0.22 ± 0.04 µm2 in L1b with only a slight variability between both sublaminae. The surface area of PSDs was 0.21 ± 0.01 µm2 in L1a and 0.22 ± 0.05 µm2 in L1b, respectively (Table 1; Source Data 2). The similar size of the PreAZ and PSD showed a nearly perfect overlap at the pre- and postsynaptic apposition zone. L1 AZs did not show a large variability in size as indicated by the low SD, CV, and variance (Table 1).

Revised version: On average, PreAZs were 0.18 ± 0.08 µm2 in surface area in L1a and 0.22 ± 0.04 µm2 in L1b with only a slight variability between both sublaminae. The surface area of PSDs was quite similar to that of PreAZs 0.21 ± 0.01 µm2 and was in L1a and 0.22 ± 0.05 µm2 in L1b, respectively (Table 1; Source Data 02, sheet 01). Thus, the similar size of the PreAZ and PSD showed a nearly perfect overlap at the pre- and postsynaptic apposition zone. L1 AZs did not show a large variability in size as indicated by the low SD, CV, and variance (Table 1).

Earlier version: Notably, our analysis revealed only a weak correlation between the surface area of SBs and that of PreAZs for L1a and L1b SBs (Figure 8A). These findings imply that the size of SBs and PreAZs is independently regulated from each other.

Revised version: Notably, our analysis revealed only a weak correlation between the surface area of SBs and that of PreAZs for L1a and L1b SBs (Figure 4A). These findings imply that the size of SBs and PreAZs is independently regulated from each other.

Added in the revised version: Measurements of the size of the synaptic cleft were performed on random EM images taken from the series using OpenCAR or online directly at the TEM using ImageSP (Fa. Tröndle, Moorenweis, Germany). Only synaptic clefts cut perpendicular to the PreAZ and PSD were included in the sample. The distance between the outer edge of the PreAZ and PSD was measured at the two lateral edges and separately at the central region of the synaptic cleft; the two values of the lateral edges were averaged for each cleft measurement. Finally, a mean ± SD was calculated for both the lateral and central region over all synaptic clefts analyzed per patient.

Earlier version: The width of the synaptic cleft (Table 1; Source Data 2) is an important measure at AZs since its size (diameter) critically determines the temporal and spatial neurotransmitter concentration. The cleft widths were slight, but non-significantly different between the two sublaminae and similar to findings in L4, L5 and L6 (Yakoubi et al. 2019a, b: Schmuhl-Giesen et al. 2022).

Revised version: The size of the synaptic cleft (Table 1; Source Data 02, sheet 03) is an important measure at AZs since its size (diameter) critically determines the temporal and spatial neurotransmitter concentration. The synaptic cleft size were slight, but non-significantly different between the two sublaminae and similar to findings in L4, L5 and L6 (Yakoubi et al. 2019a, b: Schmuhl-Giesen et al. 2022).

Organization of the pools of SVs in L1a and L1b excitatory SBs of the human TLN

SVs are the other key structure in neurotransmitter storage and release, hence they play a fundamental role in synaptic transmission and in the modulation of short- and long-term synaptic plasticity (Südhof 2002, 2012). Three different pools of SVs are functionally defined: the RRP, the RP, and the resting pool. Synaptic efficacy, strength, mode, and probability of release (Pr) are regulated by these pools (Schikorski and Stevens 2001; Silver et al. 2003; Rizzoli and Betz 2004; Saviane and Silver 2006; Schikorski 2014; Watanabe et al. 2014; Neher 2015; Vaden et al. 2019; reviewed by Rizzoli and Betz 2005; Denker and Rizzoli 2010; Chamberland and Toth 2016).

Earlier version: In general, SVs were distributed throughout the entire terminal in ∼95% of the population of SBs investigated (Figure 2, 3), the remainder showed a more cluster-like distribution. Two different types of vesicles were found: (1) Small clear SVs with an average diameter of 25.03 ± 5.27 nm (Table 1; Source Data 2) and (2) DCVs (Figure 3B, C, E; Source Data 2) with an average diameter of 57.97 ± 6.45 nm with no significant difference between both sublaminae. DCVs were intermingled with the population of SVs throughout the entire SB (Figure 3B, C, E).

Revised version: In general, SVs were distributed throughout the entire terminal in ∼95% of the population of SBs investigated (Figure 2, 3), the remainder showed a more cluster-like distribution. Two different types of vesicles were found: (1) Small clear SVs with an average diameter of 25.03 ± 5.27 nm (Table 1; Source Data 02, sheet 01) and (2) DCVs (Figure 3B, C, E; Source Data 02, sheet 01) with an average diameter of 57.97 ± 6.45 nm with no significant difference between both sublaminae. DCVs were intermingled with the population of SVs throughout the entire SB (Figure 3B, C, E).

Earlier version: DCVs play an important role in endo- and exocytosis build-up of PreAZs by releasing Piccolo and Bassoon (Schoch and Gundelfinger 2006; Murkherjee et al. 2010), or in clustering SVs at the PreAZs (Watanabe et al. 2014, Murkherjee et al. 2010). In addition, various co-transmitters, such as neuropeptides, ATP, noradrenalin, and dynorphin were identified in large DCVs (Ghijsen and Leenders 2005).

Revised version: DCVs play an important role in endo- and exocytosis, the build-up of PreAZs by releasing Piccolo and Bassoon (Schoch and Gundelfinger 2006; Murkherjee et al. 2010), or in clustering SVs at the PreAZs (Watanabe et al. 2014, Murkherjee et al. 2010). In addition, various co-transmitters, such as neuropeptides, ATP, noradrenalin, and dynorphin were identified in large DCVs (Ghijsen and Leenders 2005).

The distribution pattern of SVs made it impossible to morphologically distinguish between the three functionally defined pools of SVs, except for the so-called ‘docked’ and fused vesicles primed to the PreAZ (Rizzoli and Betz, 2004; Neher 2015; reviewed by Rizzoli and Betz 2005; Chamberland and Toth 2015).

Earlier version: Thus, a perimeter analysis (a modified Sholl-analysis) was used to determine the exact location of each SV from the PreAZ in individual terminals (for criteria to define SV pools see Rizzoloi and Betz 2005). The RRP were defined as a distance (perimeter p) from ≤10 nm and ≤20 nm from the PreAZ. This criteria for the RRP was chosen because both values are less than an SV diameter and hence may represent the RRP from which SVs could be easily and rapidly recruited. The RP maintains release during moderate stimulation and was estimated close (p60-p200 nm) to the PreAZ. All SVs further than 200 nm from the PreAZ constituted the resting pool which acts as a deposit of SVs only used by intense and/or repetitive stimulation. The average total pool of SVs was 2958.62 ± 1940.51 in L1a vs. 3903.32 ± 1852.23 in L1b with an average of 3430.97 ± 1773.77 for L1 (Table 1; Source Data 2). Remarkably, the total pool in L1b was significantly larger (p ≤ 0.01) by ∼1.3-fold when compared to L1a although a huge variability in total pool size exists in both sublaminae (Figure 3D, E) ranging from 758 to 6542 in L1a and 1507 to 15060 in L1b as indicated by the SD, IQR and variance (Table 1). SVs contributed with ∼5% (0.03 µm3) in L1a and ∼4% (0.02 µm3) in L1b to the total volume of SBs, however, with no significant difference between both sublaminae (Table 1; Source Data 2).

Revised version: Thus, a perimeter analysis (a modified Sholl-analysis) was used to determine the exact location of each SV from the PreAZ in individual terminals (for criteria to define SV pools see Rizzoli and Betz 2005). The RRP was defined as a distance (perimeter p) from ≤10 nm and ≤20 nm from the PreAZ. This criteria for the RRP was chosen because both values are less than an SV diameter and hence may represent the RRP from which SVs could be easily and rapidly recruited. The RP maintains release during moderate stimulation and was estimated close (p60-p200 nm) to the PreAZ. All SVs further than 200 nm from the PreAZ were regarded to constitute the resting pool which acts as a deposit of SVs only used by intense and/or repetitive stimulation.

Removed in the revised version: Interestingly, no correlation between the volume of SBs vs. the total pool size was found for L1a (R2: 0.2490), but a good correlation was observed for L1b (R2: 0.5700). Furthermore, no correlation was found between the PreAZ vs. the total pool size (Figure 4C; Source Data 2).

Removed in the revised version: The average RRP/PreAZ for L1was 3.60 ± 4.24 SVs at p10 nm (L1a: 5.90 ± 5.05 SVs and L1b: 1.29 ± 1.87 SVs) with a significant difference (p ≤ 0.01) by ∼4.5-fold between the two sublaminae (Table 1; Source Data 2). The RRP/PreAZ at p20 nm criterium was on average 19.05 ±17.23 SVs (L1a: 25.04 ± 21.09 SVs and L1b: 13.07 ± 13.87SVs) and thus nearly 2-fold larger for L1a. Interestingly, both RRPs in L1a and L1b were characterized by a comparably low variability as indicated by the SD, CV, and variance (Table 1) pointing to sublamina-specific differences in Pr, synaptic efficacy, strength and paired-pulse behavior at individual SBs. No correlation between the RRP at both the p10 nm and p20 nm criterium and the surface area of PreAZs was found for both sublaminae. (Figure 4D, E).

Removed in the revised version: The RP/PreAZ at 60-200 nm perimeter criterium was 390.12 ± 286.89 SVs in L1a and 535.88 ± 321.16 in L1b with no significant difference for both sublaminae (Table 1; Source Data 2). It should be noted however, that the variability of the RP is relatively large in both sublaminae as indicated by the SD, median, IQR and variance (Table 1).

Removed in the revised version: No correlation was observed between the SVs in the p60-200 nm RP and the surface area of the PreAZs (L1a, Figure 4F) and for L1b (Figure 4F; Source Data 2), respectively.

Earlier version: The resting pool of SVs contained on average 2515.02 ± 1588.99 SVs in L1a and 3277.97 ± 1480.96 SVs in L1b with a significant difference between the two sublaminae (p ≤ 0.01, Table 1; Source Data 2). The number of SVs in the resting pool in L1b was about 1.3-fold larger when compared to L1a. Again, also the resting pool showed a large variability for both sublaminae as indicated by the SD, CV, IQR and variance (Table 1).

Revised version: On average, the total pool of SVs was 2958.62 ± 1940.51 in L1a vs. 3903.32 ± 1852.23 in L1b with an average of 3430.97 ± 1773.77 for L1 (Table 1; Source Data 02). Remarkably, the total pool in L1b was significantly larger (p ≤ 0.01) by ∼1.3-fold when compared to L1a. A huge variability in total pool size exists in both sublaminae (Figure 3D, E) ranging from 758 to 6542 in L1a and 1507 to 15060 in L1b as indicated by the SD, IQR and variance (Table 1). SVs contributed with ∼5% (0.03 µm3) in L1a and ∼4% (0.02 µm3) in L1b to the total volume of SBs. However, no significant difference between both sublaminae was observed (Table 1; Source Data 02).

Added in the revised version: Interestingly, no correlation between the volume of SBs vs. the total pool size was found for L1a (R2: 0.2490), but a good correlation for L1b (R2: 0.5700). Furthermore, no correlation was found between the PreAZ vs. the total pool size (Figure 4C; Source Data 02).

Added in the revised version: The average number of SVs in the putative RRP/PreAZ for L1was 3.60 ± 4.24 SVs at p10 nm (L1a: 5.90 ± 5.05 SVs and L1b: 1.29 ± 1.87 SVs) with a significant difference (p ≤ 0.01) by ∼4.5-fold between the two sublaminae (Table 1; Source Data 02). The putative RRP/PreAZ at the p20 nm criterion was on average 19.05 ±17.23 SVs (L1a: 25.04 ± 21.09 SVs and L1b: 13.07 ± 13.87 SVs) and thus nearly 2-fold larger for L1a. Interestingly, both RRPs in L1a and L1b were characterized by a comparably low variability as indicated by the SD, CV, and variance (Table 1) that may point to sublaminae-specific differences in Pr, synaptic efficacy, strength and paired-pulse behavior at individual SBs.

Added in the revised version: No correlation between the RRP at both the p10 nm and p20 nm criterion and the surface area of PreAZs was found for both sublaminae. (Figure 4D, E).

Added in the revised version: The RP/PreAZ at the 60-200 nm perimeter criterion was 390.12 ± 286.89 SVs in L1a and 535.88 ± 321.16 in L1b with no significant difference for both sublaminae (Table 1; Source Data 02, sheet 01). It should be noted however, that the variability of the RP is relatively large in both sublaminae as indicated by the SD, median, IQR and variance (Table 1).

Earlier version: No correlation was observed between the SVs in the p200 nm and p500 nm resting pool and the surface area of the PreAZs in L1a and L1b (Figure 4G, H; Source Data 2).

Revised version: Again, no correlation was observed between the SVs in the p60-200 nm RP and the surface area of the PreAZs (L1a, Figure 4F) and for L1b (Figure 4F; Source Data 02, sheet 01), respectively.

Added in the revised version: The resting pool of SVs contained on average 2515.02 ± 1588.99 SVs in L1a and 3277.97 ± 1480.96 SVs in L1b with a significant difference between the two sublaminae (p ≤ 0.01, Table 1; Source Data 02, sheet 01). The number of SVs in the resting pool in L1b was about 1.3-fold larger when compared to L1a. Again, also the resting pool showed a large variability for both sublaminae as indicated by the SD, CV, IQR and variance (Table 1).

Added in the revised version: No correlation was observed between the SVs in the p200 nm and p500 nm resting pool and the surface area of the PreAZs in L1a and L1b (Figure 4G, H; Source Data 02, sheet 01).

Added in the revised version: Neither for the total nor for the three different pools of synaptic vesicles significant differences between SBs terminating on dendritic spines or shafts were found (Source Data 02, sheet 04).

Earlier version: EM tomography of L1 excitatory SBs in the human TLN

Revised version: TEM tomography of L1 excitatory SBs in the human TLN

Earlier version: It is still controversially discussed whether so-called ‘docked’ SVs represent the RRP. To compare our results from the perimeter analysis for the p10 nm RRP with that of ‘docked’ SVs, high-resolution EM tomography was carried out. Only SBs where the AZ could be followed from its beginning to its end in individual tilt-series and where the AZ was cut perpendicular through the PreAZ, PSD and the synaptic cleft were analyzed. In L1a (50 SBs, on shafts = 25 SBs; on spines = 25 SBs) and L1b (50 SBs on shaft = 25 SBs; on spines = 25 SBs) the number of ‘docked’ SVs was analyzed at the PreAZ (Figure 5; Table 2; Source Data 3).

Revised version: It is still controversially discussed whether so-called ‘docked’ SVs or omega-shaped bodies represent the RRP. To compare our results from the perimeter analysis for the p10 nm RRP with that of ‘docked’ SVs, high-resolution TEM tomography was carried out. Only SBs where the AZ could be followed from its beginning to its end in individual tilt-series and where the AZ was cut perpendicular through the PreAZ, PSD and the synaptic cleft were analyzed. In L1a (50 SBs, on dendritic shafts = 25 SBs; on spines = 25 SBs) and L1b (50 SBs, on dendritic shafts = 25 SBs; on spines = 25 SBs) the number of ‘docked’ SVs was analyzed at the PreAZ (Figure 5; Table 2; Source Data 03).

EM tomography of SBs in L1a and L1b in the human TLN.

A, Large SB (sb) terminating with a single AZ (arrowheads) on a terminal tuft dendritic shaft (de) in L1a. B, SB (sb) synapsing on a stubby spine (sp) with a prominent spine apparatus (framed area) and a thick dendritic segment (de) in L1b. Scale bars 1 µm. C, Large mushroom (sp1) and a small spine head (sp2) receiving input from three SBs (sb1, sb2, sb3) and a single SB (sb4) in L1b. Note the large AZs (arrowheads) and the prominent spine apparatus (red framed area) in the mushroom spine. Scale bar 0.5 µm. D, E, F, High-power images of ‘docked’ SVs (arrowheads) taken from a tilt-series through an individual PreAZ at a L1a spine SB. Note the so-called omega-shaped bodies in (E, F) pointing to the already opening and release of glutamate quanta. Scale bars 0.1 µm.

The results for L1a and L1b were two-fold: First, and in agreement with already published results (Yakoubi et al. 2019a, b; Schmuhl-Giesen et al. 2022) only in a minority (<1%) of all PreAZs analyzed, regardless of their target structures, a dendritic shaft (Figure 5A) or spine (Figure 5B, C), no ‘docked’ SVs were observed. Second, the majority of PreAZs (∼98%) contained more than 2, the most 8 (L1a) and 6 (L1b) ‘docked’ SVs. This finding strongly supports multivesicular release of ‘docked’ SVs in L1a and L1b SBs in line with findings in L4, L5 and L6 SBs (Yakoubi et al. 2019a, b; Schmuhl-Giesen et al. 2022; see also Figure 5D-F).

Earlier version: On average, 3.71 ± 1.38 ‘docked’ SVs (L1a) and 3.42 ± 1.34 ‘docked SVs’ (L1b) were found at individual PreAZs with similar values for both sublaminae. Furthermore, no significant difference was found for ‘docked’ SVs between shaft vs. spine synapses in L1 (p < 0.001). Strikingly, the number of ‘docked’ SVs at L1a PreAZs was ∼2-fold smaller but ∼3-fold larger in L1b when compared to the results of our quantitative perimeter analysis for the p10 nm criterion and were also significantly different (p < 0.001). The number of ‘docked’ SVs showed a relatively low variability as indicated by the SD, CV, variance, and skewness, but were different from the values estimated with the p10 nm perimeter analysis. There seems to be a tendency that larger PreAZs contained more ‘docked’ SVs providing a larger ‘docking’ area allowing the recruitment of more SVs.

Revised version: On average, 3.71 ± 1.38 ‘docked’ SVs (L1a) and 3.42 ± 1.34 ‘docked SVs’ (L1b) were found at individual PreAZs with similar values for both sublaminae. Furthermore, no significant difference was found for ‘docked’ SVs between shaft vs. spine synapses in L1. Strikingly, the number of ‘docked’ SVs at L1a PreAZs was ∼2-fold smaller but nearly 3-fold larger in L1b when compared to the results of our quantitative perimeter analysis for the p10 nm criterion and were also significantly different (p < 0.001).

Added in the revised version: The number of ‘docked’ SVs showed a relatively low variability as indicated by the SD, CV, variance, and skewness, but were different from the values estimated with the p10 nm perimeter analysis. There seems to be a tendency that larger PreAZs contained more ‘docked’ SVs providing a larger ‘docking’ area allowing the recruitment of more SVs.

Earlier version: In summary, a notable disparity difference was observed between values obtained for L1 using the p10 nm criterion of the perimeter- and the EM tomography analysis. Additionally, the RRP in L1 and L6 was ∼2-4-times smaller when compared to values in L4 and L5 (Yakoubi et al. 2019a, b; Schmuhl-Giesen et al. 2022) pointing towards a layer-specific regulation of the RRP.

Revised version: In summary, a notable disparity difference was observed between values obtained for L1 using the p10 nm criterion of the perimeter- and the EM tomography analysis. Additionally, the RRP in L1 and L6 was ∼2-4-times smaller when compared to values in L4 and L5 (Yakoubi et al. 2019a, b; Schmuhl-Giesen et al. 2022) pointing towards a layer-specific regulation of the RRP. Moreover, the results suggest that not only ‘docked’ SVs but also those very close to the PreAZ should be considered to belong to the RRP.

Astrocytic coverage of L1 SBs in the human TLN

Earlier version: Astrocytes, by directly interacting with synaptic complexes forming the ‘tripartite’ synapse play a pivotal role in the induction, maintenance and termination of synaptic transmission by controlling the spatial and temporal concentration of neurotransmitter quanta in the synaptic cleft (reviewed by Dellarac et al. 2018). Astrocytes and their fine processes formed a relatively dense but comparably loose network within the neuropil in L1a and L1b (Figure 6A, B). However, and in contrast to L2, L3, L4 and L5 but similar to findings in L6 of the human TLN synaptic complexes in L1 were only partially ensheated by fine astrocytic processes that physically isolate synaptic complexes from the surrounding neuropil (Figure 6A, B). In only a small fraction, fine astrocytic processes were observed to form the ‘tripartite’ synapse. Moreover, the percentage (mean ± SD) of the volumetric fraction of astrocytic processes to the total volume revealed a lower, although not significantly different, amount of astrocytic processes in L1 compared to L2-L5, but comparable to L6 (Figure 6C; Source Data 4).

Revised version: Astrocytes, by directly interacting with synaptic complexes thus forming the ‘tripartite’ synapse play a pivotal role in the induction, maintenance and termination of synaptic transmission by controlling the spatial and temporal concentration of neurotransmitter quanta in the synaptic cleft (reviewed by Dellerac et al. 2018). Astrocytic profiles were identified by their irregular stellate shape, relatively clear cytoplasm, numerous glycogen granules and bundles of intermediate filaments (for criteria see Peters et al. 1991; Ventura and Harris 1999; see also Figure 6A, B). Astrocytes and their fine processes formed a relatively dense (in L4, L5), but a comparably loose network within the neuropil in L1 (Figure 6A-C; Source Data 04, sheet 01) by contributing with ∼20% to the total human TLN volume, with the highest degree in coverage in L4 but similar to findings in L6 of the human TLN (Figure 6C; Source Data 04, sheet 01).

Astrocytic coverage of synaptic complexes in the human TLN.

Earlier version: A, EM micrograph of two adjacent SBs (sb1, sb2) terminating on two spines (sp1, sp2) in L1a. Sp2 contained a prominent spine apparatus (framed area), sb1 two DCVs (asterisks). Note that both synaptic complexes were only partially ensheathed by fine astrocytic processes (transparent green) reaching the AZs only on one side. In both synaptic complexes two ‘docked’ SVs (transparent red) at the PreAZs were found. Scale 0.25 µm. B, Stubby spine (stsp) receiving two SBs (sb1, sb2) in L1b. Here, fine astrocytic processes (transparent green) were found close to the two synaptic complexes but never reached the synaptic cleft. Note also the ‘docked’ SVs (transparent red) pointing to multivesicular release. In both images the AZs are marked by arrowheads. Scale bar 0.25 µm. C, Bar histogram showing the percentage (mean ± SD) of the volumetric fraction of astrocytic processes to the total volume in L1 to L6. Values for L6 are taken from Schmuhl-Giesen et al. 2022. The horizontal bars indicate significant differences between cortical layers *: p < 0.05 (Source Data 4).

Revised version: A, EM micrograph of two adjacent SBs (sb1, sb2) terminating on two spines (sp1, sp2) in L1a. Sp2 contained a prominent spine apparatus (framed area), sb1 two DCVs (asterisks). Note that both synaptic complexes were only partially ensheathed by fine astrocytic processes (transparent green) reaching the AZs only on one side. In both synaptic complexes two ‘docked’ SVs (transparent red) at the PreAZs were found. Scale 0.25 µm. B, Stubby spine (stsp) receiving two SBs (sb1, sb2) in L1b. Here, fine astrocytic processes (transparent green) were found close to the two synaptic complexes but never reached the synaptic cleft. Note also the ‘docked’ SVs (transparent red) pointing to multivesicular release. In both images the AZs are marked by arrowheads. Scale bar 0.25 µm. C, Bar histogram showing the percentage (mean ± SD) of the volumetric fraction of astrocytic processes to the total volume in L1 to L6 (Source Data 04, sheet 01). Values for L4-L6 are taken from Schmuhl-Giesen et al. 2022.. D, Bar histogram showing the percentage (mean ± SD) of AZs covered by fine astrocytic processes either on both sides, one side or without any coverage subdivided by their location on dendritic spines or shafts (Source Data 04, sheet 02). The horizontal bars in C and D indicate significant differences *: p < 0.05; **: p < 0.01.

Added in the revised version: In general, the majority (∼70% both in spine and shaft synapses; Figure 6D, Source Data 04, sheet 02) of synaptic complexes in L1 were without astrocytic coverage close to the AZ or only partially ensheated (∼20% both in spine and shaft synapses) by fine astrocytic processes. In only a small fraction (∼10%) of both spine and shaft synapses fine astrocytic processes were observed to form the ‘tripartite’ synapse that physically isolate synaptic complexes from the surrounding neuropil. The findings in L1 are in contrast to L4 were a higher number of ‘tripartite’synaptic complexes (spines: ∼50%; shafts: ∼20%) was found. Around 25% AZs on dendritic spines and ∼ 60% on dendritic shafts were partially covered with a huge inter-individual difference. The remainder were not covered. No significant differences were found between spine and shaft synapses (Figure 6D, Source Data 04, sheet 02).

Earlier version: Hence, it is most likely that the partial coverage or complete absence of fine astrocytic processes at the majority of synaptic complexes in L1 of the human TLN may contribute to different ‘behaviors’ of synaptic complexes, for example in the removal of ‘spilled’ horizontally diffusing neurotransmitter quanta favoring synaptic crosstalk (for more detail see Discussion).

Revised version: Hence, it is most likely that the partial coverage or complete absence of fine astrocytic processes at the majority of synaptic complexes in L1 of the human TLN may contribute to different ‘behaviors’ of synaptic complexes in synaptic transmission and in the modulation of short-term plasticity, for example in the removal of ‘spilled’ horizontally diffusing neurotransmitter quanta favoring synaptic crosstalk (for more detail see Discussion).

Discussion

This study investigated the synaptic organization of L1 in the human TLN. The quantitative 3D-models analyzed showed beside similarities significant layer-specific differences in those structural and synaptic parameters, in particular in the size of the AZ and the total pool of SVs, the RRP, RP and resting pool when compared with other already investigated layers (Yakoubi et al. 2019a, b, Schmuhl-Giesen et al. 2022). The low degree of astrocytic coverage of L1 SBs suggests that glutamate spillover and as a consequence synaptic cross talk may occur at the majority of synaptic complexes in L1

Synaptic density measurements

Earlier version: Synaptic density measurement are a useful tool to describe the synaptic organization of a particular area, nuclei and even layers in different brain regions, but also the degree of connectivity underlying the computational properties of a given brain area or in a given brain network. Meanwhile numerous studies in various animal species and brain regions have performed such an analysis but data for a density of synaptic complexes in humans are still rare (but see Marco and DeFelipe 1998; Tang et al. 2001; DeFelipe et al. 1999, 2002; Alonso-Nanclares et al. 2008; Blazquez-Llorca et al. 2013; Finnema et al. 2016; Cano-Astorga et al. 2021, 2023).

Revised version: Synaptic density measurements are a useful tool to describe the synaptic organization of a particular area, nuclei and even layers in different brain regions, but also the degree of connectivity underlying the computational properties of a given brain area or in a given brain network. Meanwhile numerous studies in various animal species and brain regions have performed such an analysis but data for a density of synaptic complexes in humans are still rare (but see Marco and DeFelipe 1998; Tang et al. 2001; DeFelipe et al. 1999, 2002; Alonso-Nanclares et al. 2008; Blazquez-Llorca et al. 2013; Finnema et al. 2016; Cano-Astorga et al. 2021, 2023).

Earlier version: Strikingly, a huge layer-specific difference in the mean density of synaptic contacts was found in the human TLN. In L1, the overall density was 5.26 ± 8.44*108 synaptic complexes/mm3 although a great variability was observed as indicated by the SD larger than the mean consistent with findings in the other layers of the human TLN L5 (3.89 *108 ± 9.12*108; Yakoubi et al. 2019a); L4 (2.37*106 ± 2.19*106; Yakoubi et al. 2019b), and L6 (4.98 ± 2.19*106; Schmuhl-Giesen et al. 2022) and existing data published by Alonso-Nanclares et al. 2008: 9.13 ± 0.63*108 and Tang et al. 2001: 164*1012. The density values for L1 and L5 are in the same order of magnitude but differed substantially for L4 and L6b by one or two orders of magnitude. In summary, the highest synaptic density was observed in L1, indicating a heightened level of connectivity and synaptic interaction. This suggests that L1 facilitates rapid information processing due to its robust synaptic activity.

Revised version: Strikingly, a huge layer-specific difference in the mean density of synaptic contacts was found in the human TLN (see also Supplemental Table 1). In L1, the overall density was 5.26*108 ± 8.44*107 synaptic complexes/mm3 although a great variability was observed as indicated by the SD consistent with findings in the other layers of the human TLN L4 (2.37*106 ± 2.19*106; Yakoubi et al. 2019b); L5 (3.89 *108 ± 9.12*108; Yakoubi R., PhD Thesis 2023) and L6 (4.98*107 ± 1.85*107; Schmuhl-Giesen et al. 2022) and existing data published by Alonso-Nanclares et al. 2008: 9.13 ± 0.63*108 and Tang et al. 2001: 164*1012. The density values for L1 and L5 are in the same order of magnitude but differ substantially from L4 and L6 by one or two orders of magnitude.

Added in the revised version: In summary, the highest synaptic density was observed in L1, indicating a heightened level of connectivity and synaptic interaction. This suggests that L1 facilitates rapid information processing due to its robust synaptic activity.

Important structural subelements of SBs in the human TLN Shape and size of PreAZs and PSDs

Added in the revised version: Synaptic efficacy, strength, modes of release and Pr are beside the pool of SVs, determined by the shape and size of AZs (Matz et al. 2010; Holderith et al. 2012; Südhof 2012).

Earlier version: Synaptic efficacy, strength, modes of release and Pr are beside the pool of SVs, determined by the shape and size of AZs (Matz et al. 2010; Holderith et al. 2012; Südhof 2012). The majority of SBs in L1 of the human TLN had a single at most three AZs that could be of the non-perforated macular or perforated type comparable with results for other layers in the human TLN (L4 Yakoubi et al. 2019b, L5: Yakoubi et al. 2019a, L6: Schmuhl-Giesen et al. 2022) but by ∼1.5-fold larger than in rodent and non-human primates (Marrone et al. 2005; Rollenhagen et al. 2015, 2018; Bopp et al. 2017; Hsu et al. 2017).

Revised version: The majority of SBs in L1 of the human TLN had a single at most three AZs that could be of the non-perforated macular or perforated type comparable with results for other layers in the human TLN (L4 Yakoubi et al. 2019b, L5: Yakoubi et al. 2019a, L6: Schmuhl-Giesen et al. 2022) but by ∼1.5-fold larger than in rodent and non-human primates (Marrone et al. 2005; Rollenhagen et al. 2015, 2018; Bopp et al. 2017; Hsu et al. 2017).

Earlier version: The surface area of AZs in L1 SBs of the human TLN was on average ∼0.20 μm2. This is in good agreement with data obtained for AZs in L3 (Cano-Astorga et al. 2021), L4 (Yakoubi et al. 2019b), L5 (Yakoubi et al. 2019a), L6 (Schmuhl-Giesen et al. 2022) and in L3 (Cano-Astorga et al. 2023) of the human temporal and cingulate neocortex. However, they were ∼2- to 3-fold larger than those in mouse and non-human primates visual, motor, and somatosensory neocortex (Bopp et al. 2017; Hsu et al. 2017) but remarkably even larger than AZs in comparably large CNS terminals like the hippocampal mossy fiber bouton (Rollenhagen et al. 2007) the cerebellar mossy fiber bouton (Xu-Frieman and Regehr 2003)and the Calyx of Held in the medial nucleus of the trapezoid body (Sätzler et al. 2002).

Revised version: The surface area of AZs in L1 SBs of the human TLN was on average ∼0.20 μm2. This is in good agreement with data obtained for AZs in L3 (Cano-Astorga et al. 2021), L4 (Yakoubi et al. 2019b), L5 (Yakoubi et al. 2019a), L6 (Schmuhl-Giesen et al. 2022) and in L3 (Cano-Astorga et al. 2023) of the human temporal and cingulate neocortex. However, they were ∼2- to 3-fold larger than those in mouse and non-human primates visual, motor, and somatosensory neocortex (Bopp et al. 2017; Hsu et al. 2017) but remarkably even larger than AZs in comparably large CNS terminals like the hippocampal mossy fiber bouton (Rollenhagen et al. 2007) the cerebellar mossy fiber bouton (Xu-Friedman and Regehr 2003)and the Calyx of Held in the medial nucleus of the trapezoid body (Sätzler et al. 2002).

The substantial variability in the size of AZs at individual SBs, observed both in the human TLN and in experimental animals, likely contributes to differences in synaptic efficacy, strength, Pr, quantal size, as well as the size of the RRP and RP (Südhof 2002; Matz et al. 2010, Freche et al. 2011; Holderith et al. 2012; Neher 2015; Chamerland and Toth 2016; Rollenhagen et al. 2018; Vaden et al. 2019; reviewed by Rizzoli and Betz 2005; Denker and Rizzoli 2010). However, it has to be noted that the size of the AZ is regulated as a function of activity, as shown for hippocampal SBs in the CA1 subregion (Matz et al. 2010; Holderith et al. 2012). However, the comparably large size of the AZs in the human TLN also suggest a larger ‘docking’ area allowing the fusion of more SVs and thus a larger Pr.

Size of the three pools of SVs

Besides the size of the PreAZ, the pool of rapidly releasable SVs also critically determines Pr and thus synaptic efficacy, strength and plasticity (Rosenmund and Stevens 1996; Schikorski and Stevens 2001; Rizzoli and Betz 2004; Schikorski 2014; Watanabe et al. 2014; Vaden et al. 2019; reviewed by Rizzoli and Betz 2005, Neher 2015; Chamberland and Toth 2016). It is still rather unclear whether functionally heterogeneous SV pools are structurally identifiable and thus support diverse forms of synaptic transmission und would also play a pivotal role in long- and short-term plasticity. Synaptic transmission can be modulated in various ways depending on the availability of SVs and on their recycling rates. Hence, the size of both the RRP and RP critically determines synaptic efficacy, strength and plasticity. These parameters are controlled at the PreAZ but vary substantially across various CNS synapses (reviewed by Rizzoli and Betz 2005; Neher 2015). The contribution of the RRP size to synaptic dynamics and the mechanisms by which such control is achieved at individual SBs remains largely unknown.

L1 SBs had a total pool size of ∼3500 SVs/AZ, the largest total pool size when compared with other layers of the human TLN by ∼2-fold (L4, ∼1800 SVs) by ∼2.6-fold (L5, ∼1350 SVs) and by ∼3-fold (L6, ∼1150 SVs), respectively. Remarkably, the total pool sizes in the human TLN were significantly larger by more than 6-fold (∼550 SVs/AZ), and ∼4.7-fold (∼750 SVs/AZ;) than those in L4 and L5 (Yakoubi et al. 2019a, b; see also Rollenhagen et al. 2018) in rats. The largest total pool size in L1 may point to a pivotal role in the rapid replenishment after depletion by high-frequency stimulation and the transfer of SVs into the RRP and RP from the total pool in L1 SBs.

Earlier version: The putative RRP at L1 PreAZs was on average ∼4 SVs/AZ for the p10 nm RRP and similar to that in L5 (∼5 SVs/AZ), but significantly smaller by ∼5.3-fold to that of L4 (∼20 SVs/AZ) and by ∼3.5-fold when compared to L6 in the human TLN. Also, huge differences were observed in the p20 nm RRP which at L1 PreAZs was ∼20 SVs, ∼40 SVs in L4, ∼15 SVs in L5 and ∼30 SVs in L6, respectively. It has to be noted that a layer-specific difference in both the p10 nm and p20 nm RRP exit although a huge variability was observed as indicated by the SD, IQR, and variance. Hence, at L1 PreAZs the overall RRP, taken the p10 nm and p20 nm criterion (both together are less than a vesicle diameter) was constituted by ∼25 SVs rapidly available SVs that may partially contribute to the efficacy, strength and reliability of synaptic transmission, but also in the modulation of short-term synaptic plasticity in L1.

Revised version: The putative RRP at L1 PreAZs was on average ∼4 SVs/AZ for the p10 nm RRP and similar to that in L5 (∼5 SVs/AZ), but significantly smaller by ∼5.3-fold to that of L4 (∼20 SVs/AZ) and by ∼3.5-fold when compared to L6 in the human TLN. Also, huge differences were observed in the p20 nm RRP which at L1 PreAZs was ∼20 SVs, ∼40 SVs in L4, ∼15 SVs in L5 and ∼30 SVs in L6, respectively. It has to be noted that a layer-specific difference in both the p10 nm and p20 nm RRP exists although a huge variability was observed as indicated by the SD, IQR, and variance. Hence, at L1 PreAZs the overall RRP, taken the p10 nm and p20 nm criterion (both together are less than a vesicle diameter) was constituted by ∼25 SVs rapidly available SVs that may partially contribute to the efficacy, strength and reliability of synaptic transmission, but also in the modulation of short-term synaptic plasticity in L1.

The size of the putative RP/PreAZ was ∼470 SVs in human L1 SBs and thus ∼3-fold larger to that in L6 SBs, ∼1.3-fold larger than in L4, and ∼2.4-fold larger than in L5 SBs of the human TLN. In the rodent neocortex the RP/PreAZ comprised ∼130 (L4; see also Rollenhagen et al. 2015) and ∼200 SVs (L5; see also Rollenhagen et al. 2018). The comparably largest size of the RP in L1 when compared with values for the other cortical layers in the human TLN and that found for L4 and L5 in rodents may also point to a rapid availability of SVs from the RP after the replenishment of the RRP during high-frequency stimulation. If the refilling rates were activity dependent, the large size of the RP at L1 PreAZs could explain some forms of short-term synaptic plasticity, e.g. a substantial increase in synaptic strength during frequency facilitation and post-tetanic potentiation at these SBs.

Finally, the resting pool of SVs at L1 PreAZs in the human TLN is also comparably large (∼3000 SVs) and again the largest when compared with other layers in the human TLN (L4 and L5: ∼1250 SVs; L6: ∼900 SVs). The size of the resting pool may guarantee to rapidly replenish the RRP and RP after repetitive high-frequency stimulation via active transfer of SVs with the help of mitochondria associated with the pool of SVs (this study, see also Zhu and Fuster 1996; Verstreken et al. 2005; Smith et al. 2016).

The notable disparities in AZ and SV pool sizes among individual SBs may contribute to rapid alterations in the computational properties of single neurons or networks. Consequently, these variations in AZ and SV pool sizes at L1 SBs may critically influence the behavior of SBs during so-called Up-and-Down states as described for other SBs of the CNS (Zhou and Fuster 1996; Sanchez-Vives and McCormick 2000; Sakata and Harris 2009; Testa-Silva et al. 2014).

Earlier version: Importance of presynaptic mitochondria for synaptic transmission

Revised version: Importance of presynaptic mitochondria in synaptic transmission

Mitochondria in the cortical layers of the human TLN were organized in clusters associated with the pool of SVs (this study, see also Yakoubi et al. 2019a, b; Schmuhl-Giesen et al. 2022). In L1, mitochondria only contribute by about 7.2% (L1a) and 6.7% (L1b) to the total SB volume with similar values between L1 and L6, but their percentage was ∼2-fold lower than values in L4 and L5 suggesting a layer-specific distribution of mitochondria in the human TLN.

Mitochondria play a role in the recruitment and mobilization of SVs from the RP and resting pool and in the priming and docking process of SVs (Verstreken et al. 2005; Perkins et al. 2010; Smith et al. 2016; reviewed by Dallerac et al. 2018). In the CNS they act as the main source of internal calcium (Pozzan and Rizzuto 2000; Rizzuto et al. 2000), thus they regulate and control the internal calcium concentrations in CNS terminals required for the signal cascades where for example, synaptic proteins driven by Ca2+ like synaptotagmin, synaptophysin, synaptobrevin and the SNARE-complex are involved (Südhof 2002).

Earlier version: Glial coverage of L1 SBS in the human TLN

Revised version: Astrocytic coverage of L1 SBS in the human TLN

Earlier version: It is widely recognized that astrocytes play a crucial role in the formation of the ‘tripartite’ synapse, which is a common characteristic of cortical synapses. Astrocytes serve as both a physical barrier to glutamate diffusion and as mediate neurotransmitter uptake via transporters, thereby regulating the spatial and temporal concentration of neurotransmitters in the synaptic cleft (Oliet et al. 2004; Min and Nevian 2012; Pannasch et al. 2014; reviewed by Dallerac et al. 2018). In addition, they modulate synaptic transmission by activating pre- and postsynaptic receptors (Haydon and Carmignoto 2006; Le Meur et al. 2012). Moreover, it was found that the control of t-LTD at neocortical synapses is critically influenced by astrocytes by increasing Ca2+ signaling during the induction of t-LTD (Min and Nevian 2012).

Revised version: It is widely recognized that astrocytes play a crucial role in the formation of the ‘tripartite’ synapse, which is a common characteristic of cortical synapses. Astrocytes serve as both a physical barrier to glutamate diffusion and as mediate neurotransmitter uptake via transporters, thereby regulating the spatial and temporal concentration of neurotransmitters in the synaptic cleft (Oliet et al. 2004; Min and Nevian 2012; Pannasch et al. 2014; reviewed by Dallerac et al. 2018). In addition, they modulate synaptic transmission by activating pre- and postsynaptic receptors (Haydon and Carmignoto 2006; Le Meur et al. 2012). Moreover, it was found that the control of timing-dependent long-term depression (t-LTD) at neocortical synapses is critically influenced by astrocytes by increasing Ca2+ signaling during the induction of t-LTD (Min and Nevian 2012).

Earlier version: Remarkably, significant layer-specific differences exist in the astrocytic coverage of synaptic complexes in the human TLN. Whereas in L4 and L5 ∼80% of the total volumetric fraction was occupied by astrocytic processes and most of synaptic complexes were tightly ensheathed by fine astrocytic processes, in L1 and L6 the volumetric fraction of astrocytic processes was only ∼30% with a loose and incomplete coverage of synaptic complexes. As a consequence, the lack or incomplete astrocytic coverage of synaptic complexes in L1 can only partially act as a physical barrier to neurotransmitter diffusion. In addition, beside a vertical also its horizontal diffusion of glutamate at the synaptic cleft is possible. Hence, two possible scenarios for the role of astrocytes are present at L1 synaptic complexes. First, as shown for only a small part of synaptic complexes in L1 the ‘tripartite’ synapse is realized at both sides of the synaptic cleft. This would allow the selective uptake of horizontally diffusing glutamate via glutamate transporters located in the fine astrocytic processes. This prevents ‘glutamate spillover’, and thus allow a vertically directed neurotransmitter diffusion and docking to postsynaptic receptors at the PSD.

Revised version: Remarkably, significant layer-specific differences exist in the astrocytic coverage of synaptic complexes in the human TLN. Whereas in L4 and L5 ∼60-80% of the total volumetric fraction was occupied by astrocytic processes and most of AZs in L4 were tightly ensheathed by fine astrocytic processes, in L1 and L6 the volumetric fraction of astrocytic processes was only ∼20% with a loose and incomplete coverage of AZs in L1 (Figure 6D). As a consequence, the lack or only incomplete astrocytic coverage of synaptic complexes in L1 can only partially act as a physical barrier to neurotransmitter diffusion. In addition, beside a vertical also its horizontal diffusion of glutamate at the synaptic cleft is possible. Hence, two possible scenarios for the role of astrocytes are present at L1 synaptic complexes. First, as shown for only a small part (10%) of synaptic complexes in L1 the ‘tripartite’ synapse is realized at both sides of the synaptic cleft (Figure 6D). This would allow the selective uptake of horizontally diffusing glutamate via glutamate transporters located in the fine astrocytic processes. This prevents ‘glutamate spillover’, and thus allow a vertically directed neurotransmitter diffusion and docking to postsynaptic receptors at the PSD.

Earlier version: The remaining larger part of synaptic complexes in L1 either lack a coverage by fine astrocytic processes or are only partially ensheathed. Here ‘glutamate spillover’ of horizontally diffusing neurotransmitter quanta is most likely and thus synaptic crosstalk between neighboring synaptic complexes at AZs in neighboring spines on terminal tuft dendrites because spine density in L1 is high. This may cause a switch from asynchronous to synchronous release from neighboring synaptic complexes upon repetitive low- and high-frequency stimulation (von Gersdorff and Borst 2002).

Revised version: The remaining larger part of synaptic complexes in L1 either lack a coverage by fine astrocytic processes at AZs or are only partially ensheathed (Figure 6D). Here ‘glutamate spillover’ of horizontally diffusing neurotransmitter quanta is most likely and thus synaptic crosstalk between neighboring synaptic complexes at AZs in neighboring spines on terminal tuft dendrites because spine density in L1 is high. This may cause a switch from asynchronous to synchronous release from neighboring synaptic complexes upon repetitive low- and high-frequency stimulation (von Gersdorff and Borst 2002).

In summary, two distinct structural scenarios were identified regarding the role of astrocytes in L1, However, it is still rather unknown which of the two scenarios contribute more efficiently to synaptic transmission and plasticity in L1.

The role of L1 in the information processing of the cortical column

Earlier version: The computational properties of the neocortex depend upon its ability to integrate the information provided by the sensory organs (bottom-up information) with internally generated signals such as expectations or attentional signals (top-down information). This integration occurs in apical tuft dendrites in L1. Importantly, L1 is the predominant input layer for top-down information, relayed by a rich, dense network of cortico-cortical, commissural and associational long-range axonal projections and L1 excitatory and inhibitory neurons providing signals to the terminal tuft dendrites of pyramidal neurons (reviewed by Schumann et al. 2021). Thus, L1 is a central locus of neocortical associations controlled by distinct types of excitation and inhibition in humans (reviewed by Hartung and Letzkus 2021).

Revised version: The computational properties of the neocortex depend upon its ability to integrate the information provided by the sensory organs (bottom-up information) with internally generated signals such as expectations or attentional signals (top-down information). This integration occurs at apical tuft dendrites of L2/3 and L5 pyramidal neurons in L1. Importantly, L1 is the predominant input layer for top-down information, relayed by a rich, dense network of cortico-cortical, commissural and associational long-range axonal projections and L1 excitatory and inhibitory neurons providing signals to the terminal tuft dendrites of pyramidal neurons (reviewed by Schumann et al. 2021). Thus, L1 is a central locus of neocortical associations controlled by distinct types of excitation and inhibition in humans (reviewed by Hartung and Letzkus 2021).

Earlier version: From L4, as the first station of intracortical information processing thalamocortical signals are transferred via intracolumnar projecting axonal collaterals via L2, L3 to L1 (Feldmeyer et al. 1999; Lübke et al. 2000) representing a convergent input layer for thalamocortical and cortico-cortical synaptic inputs. Hence, L4 excitatory neurons act as ‘feed-back amplifiers’, but also guarantee a highly reliable signal transduction to L2, L3 and L1 (Feldmeyer et al. 1999; Seeman et al. 2018).

Revised version: From L4, as the first station of intracortical information processing thalamocortical signals are transferred via intracolumnar projecting axonal collaterals via L2, L3 to L1 (Feldmeyer et al. 1999; Lübke et al. 2000) representing a convergent input layer for thalamocortical (Rodriguez-Moreno et al. 2018) and cortico-cortical synaptic inputs. Hence, L4 excitatory neurons act as ‘feed-back amplifiers’, but also guarantee a highly reliable signal transduction to L3, L2 and L1 (Feldmeyer et al. 1999; Seeman et al. 2018).

Earlier version: SBs in L1 may contact both the persistent population of CR-cells, GABAergic interneurons and terminal tuft dendrites of pyramidal neurons (Figure 7) proving feed-forward excitation to their putative partners in L1. CR-cells exhibit long-range excitation within and across cortical columns via long horizontal axonal collaterals that preferentially and heavily contact terminal tuft dendrites of pyramidal neurons (Anstötz et al. 2014). As shown in rodent neocortex, GABAergic interneurons are synaptically coupled with CR-cells providing feed-forward inhibition (Anstötz et al. 2014). Another subpopulation of GABAergic interneurons in L1 provide a highly specialized form of recurrent inhibition that link cortical L1 to L3 interneurons and L5 pyramidal neurons (Jiang et al. 2013). L1 single-bouquet cells preferentially form inhibitory connections on L2/3 interneurons that inhibit the entire dendritic-somato-axonal axis of L5 pyramidal neurons located within the same column. In contrast, L1 neurogliaform cells frequently form inhibitory and electric connections with L2/3 interneurons. These L1-L3 interneurons inhibit the distal apical dendrite of > 60% of L5 pyramidal neurons across multiple columns. Functionally, the single bouquet cell to L2/3 interneurons to L5 pyramidal neuronal circuits disinhibit, and neurogliaform cells to L2/3 interneuron to L5 pyramidal neuronal circuits inhibit the initiation of dendritic complex spikes in L5 pyramidal neurons. As dendritic complex spikes can serve coincidence detection, these cortical interneuronal circuits may be essential for salience selection (Jiang et al. 2013).

Revised version: SBs in L1 contact both the persistent subpopulation of CR-cells, various GABAergic interneurons and terminal tuft dendrites of pyramidal neurons (Figure 7) proving feed-forward excitation to their putative partners in L1. CR-cells exhibit long-range excitation that preferentially and heavily contact terminal tuft dendrites of pyramidal neurons within and across cortical columns via long horizontal axonal collaterals (Anstötz et al. 2014). As shown in rodent neocortex, GABAergic interneurons are synaptically coupled with CR-cells providing feed-forward inhibition (Anstötz et al. 2014) thus controlling the output of CR-cells.

Schematic drawing of the possible wiring of L1 in the human TLN.

Earlier version: The possible neuronal and synaptic organization of neocortical L1 in humans modified and based on findings from rat neocortex (reviewed by Schuman et al. 2021). The dendritic domain of excitatory neurons is given in blue, the axonal projection and connectivity in magenta and that of different GABAergic interneurons in red and brown, respectively. Furthermore, a subpopulation of persistent CR-cells is illustrated in the diagram.

Revised version: The possible neuronal and synaptic organization of neocortical L1 in humans modified and based on findings from rat neocortex (reviewed by Schuman et al. 2021). The dendritic domain of excitatory neurons is given in blue, the axonal projection and connectivity in magenta and that of different GABAergic interneurons in red and green, respectively. Furthermore, a subpopulation of persistent CR-cells is illustrated in the diagram. Abbreviations: α7: nicotinic receptor-expressing interneuron; CanC: canopy cell; ChC: chandelier cell; L: layer; NGFC: neurogliaform cell; PC: pyramidal cell; SST: somatostatin containing interneuron; SStC: spiny stellate cells; VIP: vasoactive intestinal peptide-expressing interneuron; VIP/Bip: vasoactive intestinal peptide–expressing bipolar interneuron; VIP/CCK: cholecystokinin- and vasoactive intestinal peptide-expressing interneuron.

Added in the revised version: In general, GABAergic inhibition plays a critical role in shaping neuronal activity in the neocortex. However, most inhibitory synapses in the neocortex are formed onto pyramidal cell dendrites suggesting they may focally regulate cellular activity (Kwon et al. 2018). Somatostatin-expressing interneurons in L1 exert compartmentalized control over postsynaptic Ca2+ signals mediated by a subset of GABAergic synapses that directly target spine heads. GABAergic inhibition thus participates in localized control of dendritic electrical and biochemical signaling (Boldog et al. 2018). Furthermore, another study has shown that a surprisingly large fraction (33.4 ± 9.3%) of terminals formed symmetrical synaptic junctions onto a subset of cortical spines that were mostly co-innervated by an asymmetrical terminal. The double-innervated spines selectively received thalamocortical afferents expressing VGLUT2 but were not innervated by intracortical synaptic contacts expressing VGLUT1 (Kubota et al. 2007). These results demonstrate that inhibitory inputs to pyramidal cell spines may preferentially reduce thalamocortical rather than intracortical synaptic transmission and are therefore positioned to selectively gate extracortical information.

Added in the revised version: Another subpopulation of GABAergic interneurons in L1 exhibit a highly specialized form of recurrent inhibition that link cortical L1 to L3 interneurons and L5 pyramidal neurons (Jiang et al. 2013). L1 single-bouquet cells preferentially form inhibitory connections on L2/3 interneurons that inhibit the entire dendritic-somato-axonal axis of L5 pyramidal neurons located within the same column. In contrast, L1 neurogliaform cells frequently form inhibitory and electric connections with L2/3 interneurons. These L1-L3 interneurons inhibit the distal apical dendrite of > 60% of L5 pyramidal neurons across multiple columns. Functionally, the single bouquet cell to L2/3 interneurons to L5 pyramidal neuronal circuits disinhibit, and neurogliaform cells to L2/3 interneuron to L5 pyramidal neuronal circuits inhibit the initiation of dendritic complex spikes in L5 pyramidal neurons. As dendritic complex spikes can serve coincidence detection, these cortical interneuronal circuits may be essential for salience selection (Jiang et al. 2013).

Removed in the revised version: Abbreviations: α7: nicotinic receptor-expressing interneuron; CanC: canopy cell; ChC: chandelier cell; L: layer; NGFC: neurogliaform cell; PC: pyramidal cell; SST: somatostatin containing interneuron; VIP: vasoactive intestinal peptide-expressing interneuron; VIP/Bip: vasoactive intestinal peptide–expressing bipolar interneuron; VIP/CCK: cholecystokinin- and vasoactive intestinal peptide-expressing interneuron.

Earlier version: Finally, L1 receive GABAergic input from Martinotti cells located in L2/3 and L5, axo-axonic Chandelier cells in L2/3 and VIP-expressing interneurons in L2 all of which have a very prominent axonal plexus in L1 (Figure 7, see also Silberberg and Markram 2007; Obermayer et al. 2018; reviewed by Hartung and Letzkus 2021; Schuman et al. 2021).

Revised version: Finally, L1 receive GABAergic input from Martinotti cells located in L2/3 and L5, axo-axonic Chandelier cells in L2/3 and VIP-expressing interneurons in L2 all of which have a very prominent axonal plexus in L1 (Figure 7, see also Silberberg and Markram 2007; Obermayer et al. 2018; reviewed by Hartung and Letzkus 2021; Schuman et al. 2021; reviewed by Huang et al. 2024).

Earlier version: In summary, these interneuron populations provide inhibition and disinhibition to pyramidal terminal tuft dendrites by either shunting or promoting the generation of Ca2+- and Na+-spikes by temporal and spatial coincidence detection. Thus, the circuitry of L1 is an important and fundamental ‘integrator’ of the neocortex sampling bottom-up information within or even across cortical columns, but also act as a ‘filter’ or ‘discriminator’ but also as an ‘amplifier’ in the information processing of the neocortex.

Revised version: In summary, these GABAergic interneuron populations provide inhibition and disinhibition to pyramidal terminal tuft dendrites by either shunting or promoting the generation of Ca2+- and Na+-spikes by temporal and spatial coincidence detection. Thus, the circuitry of L1, including the persistent subpopulation of excitatory CR-cells; is an important and fundamental ‘integrator’ of the neocortex sampling bottom-up information within or even across cortical columns, but also act as a ‘filter’ or ‘discriminator’ but also as an ‘amplifier’ in the information processing of the neocortex.

Earlier version: Hence, L1 of the neocortex can be regarded as a complex and intricate layer of the brain serving different functions, for example to gate inputs, convey expectations and context as well as mediate states of consciousness, attention, cross-modal interactions, sensory perception, and learning (Bastos et al. 2012; Heeger 2017; Zagha 2020; reviewed by Gilbert and Li 2013). There is evidence that L1 processing, particularly Ca2+ signals in the distal apical dendrite, is important for sensory perception (Xu et al. 2012). Recent studies have revealed that L1 plays a crucial role in top-down attentional processes, aiding in directing attention, based on our goals and intentions (Hartung and Letzkus 2021; Schuman et al. 2021). Additionally, L1 is implicated in the formation and consolidation of long-term memories with disruptions in L1 function associated with memory deficits. Furthermore, L1 contributes to brain plasticity, which involves the brain’s ability to change and adapt in response to experience, including the process by which synapses alter in strength and number (reviewed by Hartung and Letzkus 2021; Schuman et al. 2021).

Revised version: Hence, L1 of the neocortex can be regarded as a complex and intricate layer of the brain serving different functions, for example to gate inputs, convey expectations and context as well as mediate states of consciousness, attention, cross-modal interactions, sensory perception, and learning (Bastos et al. 2012; Heeger 2017; Zagha 2020; reviewed by Gilbert and Li 2013). There is evidence that L1 processing, particularly Ca2+ signals in the distal apical dendrite, is important for sensory perception (Xu et al. 2012).

Added in the revised version: Recent studies have revealed that L1 plays a crucial role in top-down attentional processes, aiding in directing attention, based on our goals and intentions (Hartung and Letzkus 2021; Schuman et al. 2021). Additionally, L1 is implicated in the formation and consolidation of long-term memories with disruptions in L1 function associated with memory deficits. Furthermore, L1 contributes to brain plasticity, which involves the brain’s ability to change and adapt in response to experience, including the process by which synapses alter in strength and number (reviewed by Hartung and Letzkus 2021; Schuman et al. 2021).

Overall, L1 of the neocortex is a complex and fascinating layer of the brain that is involved in a wide range of cognitive processes and brain plasticity. Since this is relayed by a network of SBs in L1, their composition and connections play a pivotal role, in columnar information processing.

Material and Methods

Earlier version: Human brain tissue sampling during epilepsy surgery have been provided by Dr. med. Dorothea Miller, PD Dr. med. Marec von Lehe, Department of Neurosurgery, Knappschaftschafts/Universitäts-Krankenhaus Bochum and were approved by the Ethical Committees of the Rheinische Friedrich-Wilhelms-University/University Hospital Bonn (Ethical votum of the Medical Faculty to Prof. Dr. med. Johannes Schramm and Prof. Dr. rer. nat. Joachim Lübke, Nr. 146/11), the University of Bochum (Ethical votum of the Medical Faculty to PD Dr. med. Marec von Lehe and Prof. Dr. rer. nat. Joachim Lübke, Reg. No. 5190-14-15; and renewed Ethical votum of the Medical Faculty to Dr. med. Dorothea Miller and Prof. Dr. rer. nat. Joachim Lübke, Reg. No. 17–6199-BR). The consent of the patients was obtained by written and signed statements, and all further experimental procedures were approved by the same Ethical Committees cited above, and the EU directive (2015/565/EC and 2015/566/EC) concerning working with human tissue used for experimental and scientific purposes.

Revised version: Human brain tissue sampling during epilepsy surgery have been provided by Dr. med. Dorothea Miller, PD Dr. med. Marec von Lehe, Department of Neurosurgery, Knappschaftschafts/Universitäts-Krankenhaus Bochum and were approved by the Ethical Committees of the Rheinische Friedrich-Wilhelms-University/University Hospital Bonn (Ethical votum of the Medical Faculty to Prof. Dr. med. Johannes Schramm and Prof. Dr. rer. nat. Joachim Lübke, Nr. 146/11), the University of Bochum (Ethical votum of the Medical Faculty to PD Dr. med. Marec von Lehe and Prof. Dr. rer. nat. Joachim Lübke, Reg. No. 5190-14-15; and renewed Ethical votum of the Medical Faculty to Dr. med. Dorothea Miller and Prof. Dr. rer. nat. Joachim Lübke, Reg. No. 17–6199-BR). The consent of the patients was obtained by written and signed statements, and all further experimental procedures were approved by the same Ethical Committees cited above, and the EU directive (2015/565/EC and 2015/566/EC) concerning working with human tissue used for experimental and scientific purposes. All subsequent experimental procedures were approved by the Research Committee of the Research Centre Jülich GmbH.

Removed in the revised version: All subsequent experimental procedures were approved by the Research Committee of the Research Centre Jülich GmbH.

Fixation and tissue processing for TEM

Earlier version: Tissue samples from the human TLN were after their removal prepared and embedded for conventional TEM and EM tomography analysis. All neocortical access tissues were obtained from patients (1 male and 3 female, 24-65 years in age, see Supplemental Table 1) suffering from drug-resistant temporal lobe epilepsy. The pre-surgical work-up comprised at least high-resolution magnetic resonance imaging together with long-term video-electro-encephalography (EEG)-monitoring. In all cases, the circumscribed epileptic focus was in the hippocampus proper, but not in the neocortical access regions of the TL.

Revised version: Tissue samples from the human TLN were after their removal prepared and embedded for conventional TEM and EM tomography analysis. All neocortical access tissues were obtained from patients (1 male and 3 female, 24-65 years in age, see Supplemental Table 2) suffering from drug-resistant temporal lobe epilepsy. The pre-surgical work-up comprised at least high-resolution magnetic resonance imaging together with long-term video-electro-encephalography (EEG)-monitoring. In all cases, the circumscribed epileptic focus was in the hippocampus proper, but not in the neocortical access regions of the TL.

Earlier version: During epilepsy surgery, blocks of both non-affected, non-epileptic and epileptic neocortical access tissue samples (see Supplemental Table 1) were resected to control the seizures for histological inspection by neuropathologists. The non-affected non-epileptic neocortical access tissue samples was always taken far from the epileptic focus and may thus be regarded as non-epileptic as also demonstrated by other structural and functional studies using the same experimental approach (Alonso-Nanclares 2008; Navarette et al. 2013; Mohan et al. 2015; Molnár et al. 2016, Seeman et al. 2018; reviewed by Mansfelder et al. 2019). So-called ‘post mortem’ tissue was not used in this study since the ultrastructural quality (preservation) of such material is not suitable enough for fine-scale high-resolution EM, due to severe distortions of relevant structural features, e.g. fragmentation and lysis of membranes of SBs, PreAZs, PSDs and SVs that are required for the generation of quantitative 3D-models of SBs and their target structures (Lübke, personal observation). However, under certain conditions also ‘post mortem’ tissue samples can be used (Dominguez-Alvaro et al. 2019, 2921; Cano-Astorga et al. 2021, 2023).

Revised version: During epilepsy surgery, blocks of both non-affected, non-epileptic and epileptic neocortical access tissue samples (see Supplemental Table 2) were resected to control the seizures for histological inspection by neuropathologists. The non-affected non-epileptic neocortical access tissue samples was always taken far from the epileptic focus and may thus be regarded as non-epileptic as also demonstrated by other structural and functional studies using the same experimental approach (Alonso-Nanclares 2008; Navarette et al. 2013; Mohan et al. 2015; Molnár et al. 2016, Seeman et al. 2018; reviewed by Mansfelder et al. 2019). So-called ‘post mortem’ tissue was not used in this study since the ultrastructural quality (preservation) of such material is not suitable enough for fine-scale high-resolution EM, due to severe distortions of relevant structural features, e.g. fragmentation and lysis of membranes of SBs, PreAZs, PSDs and SVs that are required for the generation of quantitative 3D-models of SBs and their target structures (Lübke, personal observation). However, under certain conditions also ‘post mortem’ tissue samples can be used (Dominguez-Alvaro et al. 2019, 2921; Cano-Astorga et al. 2021, 2023).

Earlier version: Inter-individual differences were found in the structural and synaptic parameters analyzed as shown by the box plots (Supplemental Figures 1, 2; Source Data 2).

Revised version: Inter-individual differences were found in the structural and synaptic parameters analyzed as shown by the box plots (Supplemental Figures 1, 2; Source Data 02).

Other recent studies using the same experimental approach neglected and thus discarded the effect of the pharmacological treatments and disease condition (Alonso-Nanclares et al. 2008; Testa Silva et al. 2014; Mohan et al. 2015; Molnár et al. 2016; Seeman et al. 2018; Dominguez-Alvaro et al. 2019; Yakoubi et al. 2019a, b; Schmuhl-Giesen et al. 2022).

For the study of L1 in the human TLN, blocks of neocortical access tissue were sampled from the temporo-lateral or temporo-basal regions of the inferior temporal gyrus and the gyrus medialis. Immediately after their removal during epilepsy surgery, biopsy samples of the TLN were immersion-fixed in ice-cold 4% paraformaldehyde and 2.5% glutaraldehyde diluted in 0.1 M phosphate buffer (PB, pH 7.4) for 24-72 hrs at 4°C. The fixative was replaced by fresh fixative after 2 hrs and changed twice during the subsequent fixation period. Prior to vibratome sectioning, brain tissue samples were thoroughly rinsed in ice-cold PB and afterwards embedded in 5% Agar-Agar (Sigma, Munich, Germany) diluted in PB.

Neocortical tissue blocks were cut in the coronal plane through the TLN with a Vibratome VT 1000S (Leica Microsystems GmbH, Wetzlar, Germany) into 150-200 µm thick sections, collected in ice-cold PB, and washed again several times in PB. Afterwards they were transferred to 0.5-1% PB-sucrose buffered Osmium tetroxide (OsO4, 300 mOsm, pH 7.4; Sigma, Munich, Germany) for 60-90 min. After visual inspection to check for the quality of post-fixation, sections were thoroughly washed several times in PB and left overnight at 4°C in PB. The next day they were dehydrated in an ascending series of ethanol starting at 20%, 30%, 50%, 60%, 70%, 80%, 90%, 95% to absolute ethanol (15 min for each step and absolute ethanol, 30 min twice), followed by a brief incubation in propylene oxide (2 min twice; Fluka, Neu-Ulm, Germany). Sections were then transferred into a mixture of propylene oxide and DurcupanTM resin (2:1, 1:1 for 1h each; Fluka, Neu-Ulm, Germany) and stored overnight in pure resin. The next day, sections were flat embedded on coated glass slides in fresh DurcupanTM, coverslipped with AclaTM foils and polymerized at 60°C for 2 days.

Semi- and ultrathin sectioning

After light microscopic (LM) inspection, a tissue block containing the region of interest (ROI) was glued on a pre-polymerized block and trimmed down. Semithin sections were cut with a Leica UltracutS ultramicrotome (Leica Microsystems, Vienna, Austria), with a Histo-Diamond knife (Fa. Diatome, Nidau, Switzerland). Afterwards they were briefly stained with methylene-blue (Sigma-Aldrich Chemie GmbH, Taufkirchen, Germany) to identify the cortical layers, particularly L1 and underlying L2, examined and documented using a motorized Olympus BX61 microscope equipped with the Olympus CellSense analysis hard- and software (Olympus GmbH, Hamburg, Germany). All images were stored in a database until further use.

After further trimming of the block to its final size, serial ultrathin sections (50 ± 5 nm thickness; silver to silver gray interference contrast) were cut with a Leica UltracutS ultramicrotome through the determined ROI of L1a and L1b, respectively. Ultrathin sections in serial sequence were collected on pioloform-coated slot copper grids (Plano, Wetzlar, Germany). Prior to EM examination, sections were stained with 5% aqueous uranyl acetate or Uranyless (Science Services, Munich, Germany) for 15–20 min and lead citrate for 3–5 min according to Reynolds (1963) to enhance the contrast of biological membranes at the EM level. In each series of ultrathin sections, two or three ROIs throughout L1a or L1b were chosen, and photographed at a final EM magnification of 8000x with a Zeiss Libra 120 (Fa. Zeiss, Oberkochen, Germany) equipped with a Proscan 2K digital camera (Fa. Tröndle, Moorenweis, Germany) using the ImageSP software and its panorama function (Fa. Tröndle, Moorenweis, Germany). In addition, interesting details of synaptic structures in L1a and L1b were taken at various EM magnifications.

All images were stored in a database until further use. Selected EM images for publication were further edited using Adobe PhotoshopTM and Adobe IllustratorTM software.

3D-volume reconstructions and quantitative analysis of SBs

EM panorama images composing each series were imported, stacked, and aligned in the reconstruction software OpenCAR (Contour Alignment Reconstruction; for details see Sätzler et al. 2002). The main goal of this study was to quantify several structural and synaptic parameters representing structural correlates of synaptic transmission and plasticity. Excitatory SBs were characterized by large round SVs and prominent PreAZs and PSDs in contrast to putative GABAergic terminals that have smaller, more oval-shaped SVs, and a thin or no PSDs.

The following structural parameters were analyzed: (1) surface area and volume of SBs; (2) volume of mitochondria; (3) surface area of PreAZs (Dufour et al. 2016) and PSDs; two apposed membrane specializations separated by the synaptic cleft; (4) number and diameter of clear SVs and dense-core vesicles (DCVs); and (5) distance of individual SVs from the PreAZ for the structural definition of the RRP, RP, and resting pool. 3D-volume reconstructions were generated by drawing contour lines on the structures of interest in each panorama image within a series of images until a given structure was completed. Presynaptic SBs, their mitochondria as well as their postsynaptic target structures were outlined on the outer edge of their membranes, using closed contour lines throughout the entire stack within a series. A SB was considered completely captured, when it was possible to follow the axon in both directions through the entire series (en passant SBs) or the enlargement of the axon leading to an endterminal SB. The beginning of a synaptic terminal was defined by the typical widening of the axon and the abrupt occurrence of a pool of SVs.

The PreAZs and PSDs were regarded as complete when their perimeters were entirely reconstructed in a series of ultrathin sections. The surface areas of the PreAZ and PSD were computed separately by first generating a 3D surface model of the SB. The PreAZ was then measured by extracting this area from the reconstructed presynaptic bouton membrane that was covered by this membrane specialization. Hence, the length (L) of the PreAZ (L PreAZ) and the surface area (SA) of the PreAZ (SA PreAZ) is already known.

Earlier version: The size of the PSD opposing the PreAZ was estimated under the following assumptions: (1) both membrane specializations, PreAZ and PSD run parallel to each other at the pre- and postsynaptic apposition zone; (2) for both membrane specializations a contour line was drawn determining their actual length (L PreAZ and L PSD). Hence, the surface area of the PSD (SA PSD) is estimated by the following equation: which is the perimeter ratio between the outlines of the PSD to that of the synaptic contact.

Revised version: The size of the PSD opposing the PreAZ was estimated under the following assumptions: (1) both membrane specializations, PreAZ and PSD run parallel to each other at the pre- and postsynaptic apposition zone; (2) for both membrane specializations a contour line was drawn determining their actual length (L PreAZ and L PSD). Hence, the surface area of the PSD (SA PSD) is estimated by the following equation:

Added in the revised version: which is the perimeter ratio between the outlines of the PSD to that of the synaptic contact.

The synaptic cleft width was measured because of its importance for the transient temporal and spatial increase of the glutamate concentration, reversible binding of glutamate to appropriate glutamate receptors and eventual up-take and diffusion of glutamate out of the synaptic cleft. To a large extent, these processes are governed by the geometry of this structure and the shape and size of the PreAZs and PSDs.

Removed in the revised version: Measurements of the width of the synaptic cleft were performed on random EM images taken from the series using OpenCAR or online directly at the EM using ImageSP (Fa. Tröndle, Moorenweis, Germany). Only synaptic clefts cut perpendicular to the PreAZ and PSD were included in the sample. The distance between the outer edge of the pre- and postsynaptic membranes was measured at the two lateral edges and separately at the central region of the synaptic cleft; the two values of the lateral edges were averaged for each cleft measurement. Finally, a mean ± SD was calculated for both the lateral and central region over all synaptic clefts analyzed per patient.

All SVs were marked throughout each SB, their diameters were individually measured and their distances to the PreAZ were automatically detected using an algorithm implemented in OpenCAR (minimal distance between each SV membrane to the contour line of the PreAZ) throughout the SB in every single image of the series. Large DCVs were only counted in the image where they appeared largest. To avoid double counts, only clear ring-like structures were counted as SVs. However, SVs might be missed in densely packed regions, because ring-like traces may partly overlap. This effect may counteract any double counts. Based on the small extent and the partially counteracting nature of this effect, the numbers of small clear vesicles reported in this study remained uncorrected (Yakoubi et al. 2019a, b, Schmuhl-Giesen et al. 2022).

In this work, aldehyde fixation was used that is thought to induce tissue shrinkage thereby biasing structural quantification (Eyre et al. 2007; Korogod et al. 2015). A direct comparison of structural parameters obtained from either aldehyde or cryo-fixed and substituted tissue samples (Korogod et al. 2015), showed differences in cortical thickness (∼16% larger in cryo-fixed material), in the volume of extracellular space (∼6-fold larger in cryo-fixed material), a slight increase in glial volume and overall density of synaptic contacts (∼14% in cryo-fixed material), but no significant differences in neuronal structures such as axons, dendrites and SV diameter.

In the structural and synaptic parameters as estimated here, no significant difference was found for SB size and other synaptic subelements such as mitochondria, active zones (AZs) and SVs when compared with other studies (Zhao et al. 2012a, b). Therefore, no corrections for shrinkage were applied and we are thus convinced that the synaptic parameters reported here are accurate and can be directly used for detailed computational models. In addition, large-scale preservation for ultrastructural analysis will therefore continue to rely on chemical fixation approaches, due to the limited preservation of the ultrastructure in cryo-fixed material as stated in Korogod and co-workers (2015).

Golgi-Cox impregnation of biopsy material in L4 of the human TLN

Earlier version: To get an impression about the neuronal organization of the human TLN, four tissue blocks were processed with the Golgi-Cox impregnation technique using the commercially available Hito Golgi-Cox OptimStain kit (Hitobiotec Corp, Kingsport, TE, USA). After removal of the biopsy samples, tissues were briefly rinsed twice in double distilled water (dd H2O), and then transferred into impregnation solution 1 overnight at room temperature. The next day, tissue samples were incubated in fresh impregnation solution 1 and stored for 14 days in the dark at room temperature. Sections were then transferred in solution 3 and kept in the dark at room temperature for one day. Thereafter, sections were placed into fresh solution 3 in the dark at room temperature for 6 additional days. Then, solution 3 was exchanged and samples were stored at 4°C in the dark overnight. Tissue blocks were embedded in 5% agarose (Carl Roth, Karlsruhe, Germany) diluted in ddH2O, and sectioned with a vibratome in the coronal plane at 100–250 µm thickness and then transferred to ddH2O.

Revised version: To get an impression about the neuronal organization of the human TLN, four tissue blocks were processed with the Golgi-Cox impregnation technique using the commercially available Hito Golgi-Cox OptimStain kit (Hitobiotec Corp, Kingsport, TE, USA). After removal of the biopsy samples, tissues were briefly rinsed twice in double distilled water (dd H2O), and then transferred into impregnation solution 1 overnight at room temperature. The next day, tissue samples were incubated in fresh impregnation solution 1 and stored for 14 days in the dark at room temperature. The samples blocks were then transferred in solution 3 and kept in the dark at room temperature for one day. Thereafter, they were placed into fresh solution 3 in the dark at room temperature for 6 additional days. Then, solution 3 was exchanged and samples were stored at 4°C in the dark overnight. Tissue blocks were embedded in 5% agarose (Carl Roth, Karlsruhe, Germany) diluted in ddH2O, and sectioned with a vibratome in the coronal plane at 100–250 µm thickness and then transferred to ddH2O.

After careful removal of the agarose, free-floating sections were incubated into solution 3 for 2–3 min in the dark at room temperature, and right after placed into ddH2O, washed several times and stored overnight. Afterwards, they were rinsed twice in ddH2O for 4 min each, and dehydrated in 50%, 70% and 95% ethanol for 5 min each, then transferred into absolute ethanol (3 × 5 min), defatted in xylene, embedded in EukittTM (Sigma-Aldrich Chemie GmbH, Taufkirchen, Germany), finally coverslipped and air-dried. Afterwards, sections were examined and imaged with an Olympus BX 61 LM equipped with the CellSense software package (Olympus, Hamburg, Germany) at various magnifications and images were stored in a database until further use.

Stereological estimation of the density of L1 synaptic contacts in the human TLN

Earlier version: The density of synaptic complexes, composed either between an SB with a dendrite or spine, in a given volume is a valuable parameter to assess the structural and functional changes in the brain, which are linked to the age, pathological or experimental conditions (Rakic et al. 1994; DeFelipe et al. 1999). The density of synaptic contacts was unbiasedly estimated in L1, separated for L1a and L1b, from four patients, respectively (Supplemental Table 1; Source Data 1) using the physical dissector technique (Mayhew 1996; Fiala and Harris 2001) by counting the synaptic complexes in a virtual volume generated by two adjacent ultrathin sections that is the dissector: the reference section and the look-up section. Here, counting was performed using FIJI (Schindelin et al. 2012) on a stack of 20 aligned serial electron micrographs for each patient taken from the series of ultrathin sections used for the 3D-volume reconstructions of SBs in L1. An unbiased counting frame was first set and synaptic contacts to be considered (counted) are the one present in the reference section only and meeting the following criteria: presence of a PreAZ and a prominent or thin PSD separated by a synaptic cleft and SVs in the presynaptic terminal. Care was taken to distinguish between excitatory and inhibitory synaptic contacts, as well as the postsynaptic target structures (dendritic shafts or spines). Finally, the density of synaptic contacts (Nv) per 1 mm3 was calculated using the formula below: where Qd is the number of synaptic contacts per dissector and Vd is the volume of the dissector given by: Number of dissectors x frame area x section thickness.

Revised version: The density of synaptic complexes, composed either between an SB with a dendrite or spine, in a given volume is a valuable parameter to assess the structural and functional changes in the brain, which are linked to the age, pathological or experimental conditions (Rakic et al. 1994; DeFelipe et al. 1999). The density of synaptic contacts was unbiasedly estimated in L1, separated for L1a and L1b, from four patients, respectively (Supplemental Table 2; Source Data 01) using the physical dissector technique (Mayhew 1996; Fiala and Harris 2001) by counting the synaptic complexes in a virtual volume generated by two adjacent ultrathin sections that is the dissector: the reference section and the look-up section. Here, counting was performed using FIJI (Schindelin et al. 2012) on a stack of 20 aligned serial electron micrographs for each patient taken from the series of ultrathin sections used for the 3D-volume reconstructions of SBs in L1. An unbiased counting frame was first set and synaptic contacts to be considered (counted) are the one present in the reference section only and meeting the following criteria: presence of a PreAZ and a prominent or thin PSD separated by a synaptic cleft and SVs in the presynaptic terminal. Care was taken to distinguish between excitatory and inhibitory synaptic contacts, as well as the postsynaptic target structures (dendritic shafts or spines). Finally, the density of synaptic contacts (Nv) per 1 mm3 was calculated using the formula below:

Added in the revised version: where Qd is the number of synaptic contacts per dissector and Vd is the volume of the dissector given by: Number of dissectors x frame area x section thickness.

Earlier version: EM tomography of L1 SBs in the human TLN

Revised version: TEM tomography of L1 SBs in the human TLN

Earlier version: EM tomography was performed on 200-300 nm thick sections cut from blocks prepared for serial ultrathin sectioning as described above (Table 2; Source Data 3). Sections were mounted on pioloform-coated line copper grids and were counterstained with uranyl acetate and lead citrate following a slightly modified staining protocol as described by Reynolds (1063).

Revised version: EM tomography was performed on 200-300 nm thick sections cut from blocks prepared for serial ultrathin sectioning as described above (Table 2; Source Data 03). Sections were mounted on pioloform-coated line copper grids and were counterstained with uranyl acetate and lead citrate following a slightly modified staining protocol as described by Reynolds (1963).

Subsequently, sections were examined with a JEOL JEM 1400Plus, operating at 120 kV and equipped with a 4096×4096 pixels CMOS camera (TemCam-F416, TVIPS, Gauting, Germany). Tilt-series were acquired automatically over an angular range of −60° to +60° at 1°-degree increments using Serial EM (Mastronarde 2005). Stack alignment and reconstruction by filtered backprojection were carried out using the software package iMOD (Kremer et al. 1963). Final reconstructions were ultimately filtered using a median filter with a window size of 3 pixels. Tilt-series were stored as .tif files and were further processed using the freely available software Imod 4.9.12 and ImageJ (ImageJ, RRID:SCR-0033070). In each tilt-series so-called ‘docked’ SVs identified by their fusion with the PreAZ or as omega-shaped body that already released a quantum of neurotransmitter were counted separately for L1a and L1b and for their target structures, dendritic shafts or spines.

Quantitative analysis of the astrocytic coverage

Earlier version: To quantify the astrocytic coverage of synaptic complexes that may or may not constitute the ‘tripartite’ synapse in the human neocortex, the interactive software ImageJ (Schneider et al. 2012) was used. The first, the middle, and the last images of an individual EM series were used for a further quantitative volumetric analysis (Source Data 4). In each section of the same series used for the 3D-volume reconstructions, a grid (grid size 1 × 1 μm2) was placed over the EM image, and in each square, the abundance of fine astrocytic processes was documented throughout these images and averaged. Using the Cavalieri method [Unbiased Stereology: Three-Dimensional Measurement in Microscopy (Advanced Methods) Paperback-January 7, 2005, by Vyvyan Howard Matthew Reed], the (absolute) volume contribution of astrocytic processes was determined according to the Cavalieri estimator: where a(p) is the size of one square (0.8 × 0.8 μm2), P is the number of squares counted, and t is the thickness of the slice.

Revised version: Astrocytic processes around synaptic complexes were mainly identified by their relatively clear cytoplasm compared to synaptic boutons or dendritic profiles; in addition by their irregular, stellate shape, by the presence of glycogen granules and bundles of intermediate filaments (Peters et al. 1991; Ventura and Harris 1999). To quantify the astrocytic coverage of synaptic complexes that may or may not constitute the ‘tripartite’ synapse in the human neocortex, the interactive software ImageJ (Schneider et al. 2012) was used. The first, the middle, and the last images of an individual EM series were used for a further quantitative volumetric analysis (Source Data 04, sheet 01). In each section of the same series used for the 3D-volume reconstructions, a grid (grid size 1 × 1 μm2) was placed over the EM image, and in each square, the abundance of fine astrocytic processes was documented throughout these images and averaged. Using the Cavalieri method [Unbiased Stereology: Three-Dimensional Measurement in Microscopy (Advanced Methods) Paperback-January 7, 2005, by Vyvyan Howard Matthew Reed], the (absolute) volume contribution of astrocytic processes was determined according to the Cavalieri estimator:

Added in the revised version: where a(p) is the size of one square (0.8 × 0.8 μm2), P is the number of squares counted, and t is the thickness of the slice.

Added in the revised version: In addition, in consecutive sections spanning the complete AZ, the astrocytic coverage was evaluated under the following criteria: fine astrocytic processes were observed on 1) both sides of the AZ reaching as far as the synaptic cleft, 2) only one side of the AZ was covered by fine astrocytic processes and 3) AZs without astrocytic processes (Source Data 04, sheet 02). Values obtained for L1 either for SBs terminating on dendritic spines (n = 254) or shafts (n = 92) were compared to L4 (spines = 93; shafts = 41). Layer 4 was chosen because the volumetric analysis showed the highest portion of astrocytic processes in this layer compared to all other cortical layers. All values are given as percentages.

Statistical analysis

Earlier version: The mean value ± standard deviation (SD), the median with the interquartile range (IQR), the coefficient of variation (CV), skewness, variance, and the coefficient of correlation (R2) were given for each structural parameter analyzed. The p-value was considered significant only if p < 0.05. Box- and Violin plots (Plotly 4.0.0 https://chart-studio.plotly.com) were generated to investigate inter-individual differences for each patient and structural parameter (Supplemental Figure 1, 2; Source Data 2).

Revised version: The mean value ± standard deviation (SD), the median with the interquartile range (IQR), the coefficient of variation (CV), skewness, variance, and the coefficient of correlation (R2) were given for each structural parameter analyzed. The p-value was considered significant only if p < 0.05. Box- and Violin plots (Plotly 4.0.0 https://chart-studio.plotly.com) were generated to investigate inter-individual differences for each patient and structural parameter (Supplemental Figure 1, 2; Source Data 02).

Earlier version: To test for significant differences between L1a and L1b, the non-parametric Kruskal-Wallis H-test with a subsequent Mann-Whitney u-test analysis was performed, using PAST 4.02 (Hammer et al. 2001). Correlation graphs between several structural parameters were then generated (Figure 4). The R2-values were interpreted as follows: 0, no linear correlation; 0-0.5, weak linear correlation; 0.5-0.8, good linear correlation; and 0.8-1.0, strong linear correlation. Furthermore, a freely available Fisher’s r-to-z-transformation calculator (Fisher’s z) was used to test for differences in R2 between L1a and L1b (p-value < 0.05).

Revised version: To test for significant differences, the non-parametric Kruskal-Wallis H-test with a subsequent Mann-Whitney u-test analysis was performed, using PAST 4.02 (Hammer et al. 2001). Correlation graphs between several structural parameters were then generated (Figure 4). The R2-values were interpreted as follows: 0, no linear correlation; 0-0.5, weak linear correlation; 0.5-0.8, good linear correlation; and 0.8-1.0, strong linear correlation. Furthermore, a freely available Fisher’s r-to-z-transformation calculator (Fisher’s z) was used to test for differences in R2 between L1a and L1b (p-value < 0.05).

Added in the revised version: Parts of the data acquisition were carried out in the context of a PhD Thesis (Sadeghi Dastjerdi A., 2024).Values obtained in the present study for L1 were compared partially to those obtained for L2, L3 (Girwert Ch., PhD Thesis in preparation), L4, L5 (Yakoubi et al. 2019a, b; Yakoubi R., PhD Thesis 2023) and L6 (Schmuhl-Giesen et al. 2022).